Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 599 next slice time: 3
lrs-1010_7_bs=unit_only:bsr=unit_only:cond=fast:drc=off:ep=on:flr=on:lcm=predicate:nwc=1.7:stl=30:sgo=on:spo=on:spl=sat:ssfq=1.1:ssnc=none:sp=occurrence_1 on Subst-1
% (8019)Time limit reached!
% ------------------------------
% Version: Vampire 3.0 (commit e3e1a68)
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 620
% Passive clauses: 5033
% Generated clauses: 19875
% Final active clauses: 420
% Final passive clauses: 3337
% Input formulas: 29
% Initial clauses: 55
% 
% Duplicate literals: 425
% Fw subsumption resolutions: 167
% Bw subsumption resolutions: 80
% Fw demodulations: 3424
% Bw demodulations: 701
% Condensations: 187
% 
% Simple tautologies: 24
% Equational tautologies: 10
% Forward subsumptions: 10399
% Backward subsumptions: 390
% Fw demodulations to eq. taut.: 137
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 416
% Forward superposition: 8290
% Backward superposition: 4909
% Self superposition: 46
% 
% Split clauses: 32
% Split components: 66
% SAT solver clauses: 102
% SAT solver unit clauses: 32
% SAT solver binary clauses: 45
% 
% Sat splits: 32
% Sat splitting refutations: 70
% 
% Memory used [KB]: 17910
% Time elapsed: 0.400 s
% ------------------------------
% ------------------------------
% remaining time: 595 next slice time: 3
dis-1002_6_bs=off:drc=off:nwc=1.1:nicw=on:sd=1:ss=axioms:st=3.0:sio=off:spl=sat:sser=off_1 on Subst-1
% Refutation found. Thanks to Tanya!
% SZS status Theorem for Subst-1
% SZS output start Proof for Subst-1
10445. $false (0:0) [sat splitting refutation 10444,10337,10338,4914,4928,4926,4927,5163,5267]
5267. $false {2} (1:0) [trivial inequality removal 5266]
5266. sK6 != sK6 {2} (1:3) [superposition 4927,98]
98. var(sK1) = sK6 (0:4) [cnf transformation 89]
89. var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) [skolemisation 60]
60. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 59]
59. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 31]
31. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
4927. var(X50) != sK6 {2} (0:4) [sat splitting component]
5163. $false {4} (5:0) [subsumption resolution 4990,4928]
4928. sK6 = X49 {4} (0:3) [sat splitting component]
4990. subst(X28,X29,subst(X30,X31,subst(X32,X33,subst(X24,X25,app(X26,X26))))) != sK6 {4} (5:17) [backward demodulation 4928,2860]
2860. var(X27) != subst(X28,X29,subst(X30,X31,subst(X32,X33,subst(X24,X25,app(X26,X26))))) (5:18) [superposition 1094,220]
220. subst(X0,X1,app(X2,X2)) = subst(sK1,subst(X0,X1,X2),app(sK6,sK6)) (4:16) [superposition 207,145]
145. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X5,X6,app(X4,X7)) (0:16) [equality resolution 144]
144. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X6,app(X4,X7)) | X0 != X5 (0:19) [equality resolution 143]
143. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,app(X4,X7)) | X1 != X6 | X0 != X5 (0:22) [equality resolution 142]
142. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,X2) | app(X4,X7) != X2 | X1 != X6 | X0 != X5 (0:25) [equality resolution 105]
105. app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3 | subst(X0,X1,X2) != X3 | app(X4,X7) != X2 | X1 != X6 | X0 != X5 (0:28) [cnf transformation 66]
66. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3 | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [flattening 65]
65. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3) | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [ennf transformation 38]
38. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [true and false elimination 37]
37. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => ($true => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3)) [rectify 26]
26. ! [X62,X63,X64,X65,X53,X40,X51,X54] : ((X40 = X62 & X51 = X63 & app(X53,X54) = X64 & subst(X62,X63,X64) = X65) => ($true => app(subst(X40,X51,X53),subst(X40,X51,X54)) = X65)) [input]
207. app(X0,X0) = subst(sK1,X0,app(sK6,sK6)) (3:10) [superposition 192,167]
167. subst(sK1,X0,sK6) = X0 (1:6) [superposition 154,98]
98. var(sK1) = sK6 (0:4) [cnf transformation 89]
89. var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) [skolemisation 60]
60. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 59]
59. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 31]
31. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
154. subst(X4,X6,var(X4)) = X6 (0:7) [equality resolution 153]
153. subst(X0,X6,var(X4)) = X6 | X0 != X4 (0:10) [equality resolution 152]
152. subst(X0,X1,var(X4)) = X6 | X1 != X6 | X0 != X4 (0:13) [equality resolution 151]
151. subst(X0,X1,X2) = X6 | var(X4) != X2 | X1 != X6 | X0 != X4 (0:16) [equality resolution 150]
150. X5 = X6 | subst(X0,X1,X2) != X5 | var(X4) != X2 | X1 != X6 | X0 != X4 (0:19) [equality resolution 107]
107. X5 = X6 | X3 != X4 | subst(X0,X1,X2) != X5 | var(X4) != X2 | X1 != X6 | X0 != X3 (0:22) [cnf transformation 70]
70. ! [X0,X1,X2,X3,X4,X5,X6] : (X0 != X3 | X1 != X6 | var(X4) != X2 | subst(X0,X1,X2) != X5 | X3 != X4 | X5 = X6) [flattening 69]
69. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 != X3 | X1 != X6 | var(X4) != X2 | subst(X0,X1,X2) != X5) | (X3 != X4 | X5 = X6)) [ennf transformation 43]
43. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 = X3 & X1 = X6 & var(X4) = X2 & subst(X0,X1,X2) = X5) => (X3 = X4 => X5 = X6)) [true and false elimination 42]
42. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 = X3 & X1 = X6 & var(X4) = X2 & subst(X0,X1,X2) = X5) => ((X3 = X4 & $true) => X5 = X6)) [rectify 24]
24. ! [X58,X59,X60,X40,X47,X61,X51] : ((X40 = X58 & X51 = X59 & var(X47) = X60 & subst(X58,X59,X60) = X61) => ((X40 = X47 & $true) => X51 = X61)) [input]
192. app(X15,subst(sK1,X15,X16)) = subst(sK1,X15,app(sK6,X16)) (2:13) [superposition 145,167]
167. subst(sK1,X0,sK6) = X0 (1:6) [superposition 154,98]
98. var(sK1) = sK6 (0:4) [cnf transformation 89]
89. var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) [skolemisation 60]
60. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 59]
59. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 31]
31. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
154. subst(X4,X6,var(X4)) = X6 (0:7) [equality resolution 153]
153. subst(X0,X6,var(X4)) = X6 | X0 != X4 (0:10) [equality resolution 152]
152. subst(X0,X1,var(X4)) = X6 | X1 != X6 | X0 != X4 (0:13) [equality resolution 151]
151. subst(X0,X1,X2) = X6 | var(X4) != X2 | X1 != X6 | X0 != X4 (0:16) [equality resolution 150]
150. X5 = X6 | subst(X0,X1,X2) != X5 | var(X4) != X2 | X1 != X6 | X0 != X4 (0:19) [equality resolution 107]
107. X5 = X6 | X3 != X4 | subst(X0,X1,X2) != X5 | var(X4) != X2 | X1 != X6 | X0 != X3 (0:22) [cnf transformation 70]
70. ! [X0,X1,X2,X3,X4,X5,X6] : (X0 != X3 | X1 != X6 | var(X4) != X2 | subst(X0,X1,X2) != X5 | X3 != X4 | X5 = X6) [flattening 69]
69. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 != X3 | X1 != X6 | var(X4) != X2 | subst(X0,X1,X2) != X5) | (X3 != X4 | X5 = X6)) [ennf transformation 43]
43. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 = X3 & X1 = X6 & var(X4) = X2 & subst(X0,X1,X2) = X5) => (X3 = X4 => X5 = X6)) [true and false elimination 42]
42. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 = X3 & X1 = X6 & var(X4) = X2 & subst(X0,X1,X2) = X5) => ((X3 = X4 & $true) => X5 = X6)) [rectify 24]
24. ! [X58,X59,X60,X40,X47,X61,X51] : ((X40 = X58 & X51 = X59 & var(X47) = X60 & subst(X58,X59,X60) = X61) => ((X40 = X47 & $true) => X51 = X61)) [input]
145. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X5,X6,app(X4,X7)) (0:16) [equality resolution 144]
144. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X6,app(X4,X7)) | X0 != X5 (0:19) [equality resolution 143]
143. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,app(X4,X7)) | X1 != X6 | X0 != X5 (0:22) [equality resolution 142]
142. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,X2) | app(X4,X7) != X2 | X1 != X6 | X0 != X5 (0:25) [equality resolution 105]
105. app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3 | subst(X0,X1,X2) != X3 | app(X4,X7) != X2 | X1 != X6 | X0 != X5 (0:28) [cnf transformation 66]
66. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3 | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [flattening 65]
65. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3) | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [ennf transformation 38]
38. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [true and false elimination 37]
37. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => ($true => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3)) [rectify 26]
26. ! [X62,X63,X64,X65,X53,X40,X51,X54] : ((X40 = X62 & X51 = X63 & app(X53,X54) = X64 & subst(X62,X63,X64) = X65) => ($true => app(subst(X40,X51,X53),subst(X40,X51,X54)) = X65)) [input]
1094. var(X10) != subst(X11,X12,subst(X13,X14,subst(X15,X16,subst(sK1,X9,app(sK6,sK6))))) (4:18) [superposition 263,207]
207. app(X0,X0) = subst(sK1,X0,app(sK6,sK6)) (3:10) [superposition 192,167]
167. subst(sK1,X0,sK6) = X0 (1:6) [superposition 154,98]
98. var(sK1) = sK6 (0:4) [cnf transformation 89]
89. var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) [skolemisation 60]
60. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 59]
59. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 31]
31. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
154. subst(X4,X6,var(X4)) = X6 (0:7) [equality resolution 153]
153. subst(X0,X6,var(X4)) = X6 | X0 != X4 (0:10) [equality resolution 152]
152. subst(X0,X1,var(X4)) = X6 | X1 != X6 | X0 != X4 (0:13) [equality resolution 151]
151. subst(X0,X1,X2) = X6 | var(X4) != X2 | X1 != X6 | X0 != X4 (0:16) [equality resolution 150]
150. X5 = X6 | subst(X0,X1,X2) != X5 | var(X4) != X2 | X1 != X6 | X0 != X4 (0:19) [equality resolution 107]
107. X5 = X6 | X3 != X4 | subst(X0,X1,X2) != X5 | var(X4) != X2 | X1 != X6 | X0 != X3 (0:22) [cnf transformation 70]
70. ! [X0,X1,X2,X3,X4,X5,X6] : (X0 != X3 | X1 != X6 | var(X4) != X2 | subst(X0,X1,X2) != X5 | X3 != X4 | X5 = X6) [flattening 69]
69. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 != X3 | X1 != X6 | var(X4) != X2 | subst(X0,X1,X2) != X5) | (X3 != X4 | X5 = X6)) [ennf transformation 43]
43. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 = X3 & X1 = X6 & var(X4) = X2 & subst(X0,X1,X2) = X5) => (X3 = X4 => X5 = X6)) [true and false elimination 42]
42. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 = X3 & X1 = X6 & var(X4) = X2 & subst(X0,X1,X2) = X5) => ((X3 = X4 & $true) => X5 = X6)) [rectify 24]
24. ! [X58,X59,X60,X40,X47,X61,X51] : ((X40 = X58 & X51 = X59 & var(X47) = X60 & subst(X58,X59,X60) = X61) => ((X40 = X47 & $true) => X51 = X61)) [input]
192. app(X15,subst(sK1,X15,X16)) = subst(sK1,X15,app(sK6,X16)) (2:13) [superposition 145,167]
167. subst(sK1,X0,sK6) = X0 (1:6) [superposition 154,98]
98. var(sK1) = sK6 (0:4) [cnf transformation 89]
89. var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) [skolemisation 60]
60. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 59]
59. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 31]
31. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
154. subst(X4,X6,var(X4)) = X6 (0:7) [equality resolution 153]
153. subst(X0,X6,var(X4)) = X6 | X0 != X4 (0:10) [equality resolution 152]
152. subst(X0,X1,var(X4)) = X6 | X1 != X6 | X0 != X4 (0:13) [equality resolution 151]
151. subst(X0,X1,X2) = X6 | var(X4) != X2 | X1 != X6 | X0 != X4 (0:16) [equality resolution 150]
150. X5 = X6 | subst(X0,X1,X2) != X5 | var(X4) != X2 | X1 != X6 | X0 != X4 (0:19) [equality resolution 107]
107. X5 = X6 | X3 != X4 | subst(X0,X1,X2) != X5 | var(X4) != X2 | X1 != X6 | X0 != X3 (0:22) [cnf transformation 70]
70. ! [X0,X1,X2,X3,X4,X5,X6] : (X0 != X3 | X1 != X6 | var(X4) != X2 | subst(X0,X1,X2) != X5 | X3 != X4 | X5 = X6) [flattening 69]
69. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 != X3 | X1 != X6 | var(X4) != X2 | subst(X0,X1,X2) != X5) | (X3 != X4 | X5 = X6)) [ennf transformation 43]
43. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 = X3 & X1 = X6 & var(X4) = X2 & subst(X0,X1,X2) = X5) => (X3 = X4 => X5 = X6)) [true and false elimination 42]
42. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 = X3 & X1 = X6 & var(X4) = X2 & subst(X0,X1,X2) = X5) => ((X3 = X4 & $true) => X5 = X6)) [rectify 24]
24. ! [X58,X59,X60,X40,X47,X61,X51] : ((X40 = X58 & X51 = X59 & var(X47) = X60 & subst(X58,X59,X60) = X61) => ((X40 = X47 & $true) => X51 = X61)) [input]
145. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X5,X6,app(X4,X7)) (0:16) [equality resolution 144]
144. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X6,app(X4,X7)) | X0 != X5 (0:19) [equality resolution 143]
143. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,app(X4,X7)) | X1 != X6 | X0 != X5 (0:22) [equality resolution 142]
142. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,X2) | app(X4,X7) != X2 | X1 != X6 | X0 != X5 (0:25) [equality resolution 105]
105. app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3 | subst(X0,X1,X2) != X3 | app(X4,X7) != X2 | X1 != X6 | X0 != X5 (0:28) [cnf transformation 66]
66. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3 | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [flattening 65]
65. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3) | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [ennf transformation 38]
38. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [true and false elimination 37]
37. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => ($true => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3)) [rectify 26]
26. ! [X62,X63,X64,X65,X53,X40,X51,X54] : ((X40 = X62 & X51 = X63 & app(X53,X54) = X64 & subst(X62,X63,X64) = X65) => ($true => app(subst(X40,X51,X53),subst(X40,X51,X54)) = X65)) [input]
263. var(X4) != subst(X5,X6,subst(X7,X8,subst(X0,X1,app(X2,X3)))) (3:15) [superposition 206,145]
145. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X5,X6,app(X4,X7)) (0:16) [equality resolution 144]
144. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X6,app(X4,X7)) | X0 != X5 (0:19) [equality resolution 143]
143. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,app(X4,X7)) | X1 != X6 | X0 != X5 (0:22) [equality resolution 142]
142. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,X2) | app(X4,X7) != X2 | X1 != X6 | X0 != X5 (0:25) [equality resolution 105]
105. app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3 | subst(X0,X1,X2) != X3 | app(X4,X7) != X2 | X1 != X6 | X0 != X5 (0:28) [cnf transformation 66]
66. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3 | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [flattening 65]
65. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3) | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [ennf transformation 38]
38. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [true and false elimination 37]
37. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => ($true => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3)) [rectify 26]
26. ! [X62,X63,X64,X65,X53,X40,X51,X54] : ((X40 = X62 & X51 = X63 & app(X53,X54) = X64 & subst(X62,X63,X64) = X65) => ($true => app(subst(X40,X51,X53),subst(X40,X51,X54)) = X65)) [input]
206. var(X4) != subst(X5,X6,subst(X0,X1,app(X2,X3))) (2:12) [superposition 199,145]
145. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X5,X6,app(X4,X7)) (0:16) [equality resolution 144]
144. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X6,app(X4,X7)) | X0 != X5 (0:19) [equality resolution 143]
143. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,app(X4,X7)) | X1 != X6 | X0 != X5 (0:22) [equality resolution 142]
142. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,X2) | app(X4,X7) != X2 | X1 != X6 | X0 != X5 (0:25) [equality resolution 105]
105. app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3 | subst(X0,X1,X2) != X3 | app(X4,X7) != X2 | X1 != X6 | X0 != X5 (0:28) [cnf transformation 66]
66. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3 | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [flattening 65]
65. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3) | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [ennf transformation 38]
38. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [true and false elimination 37]
37. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => ($true => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3)) [rectify 26]
26. ! [X62,X63,X64,X65,X53,X40,X51,X54] : ((X40 = X62 & X51 = X63 & app(X53,X54) = X64 & subst(X62,X63,X64) = X65) => ($true => app(subst(X40,X51,X53),subst(X40,X51,X54)) = X65)) [input]
199. var(X11) != subst(X7,X8,app(X9,X10)) (1:9) [superposition 130,145]
145. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X5,X6,app(X4,X7)) (0:16) [equality resolution 144]
144. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X6,app(X4,X7)) | X0 != X5 (0:19) [equality resolution 143]
143. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,app(X4,X7)) | X1 != X6 | X0 != X5 (0:22) [equality resolution 142]
142. app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,X2) | app(X4,X7) != X2 | X1 != X6 | X0 != X5 (0:25) [equality resolution 105]
105. app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3 | subst(X0,X1,X2) != X3 | app(X4,X7) != X2 | X1 != X6 | X0 != X5 (0:28) [cnf transformation 66]
66. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3 | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [flattening 65]
65. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3) | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [ennf transformation 38]
38. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3) [true and false elimination 37]
37. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => ($true => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3)) [rectify 26]
26. ! [X62,X63,X64,X65,X53,X40,X51,X54] : ((X40 = X62 & X51 = X63 & app(X53,X54) = X64 & subst(X62,X63,X64) = X65) => ($true => app(subst(X40,X51,X53),subst(X40,X51,X54)) = X65)) [input]
130. var(X0) != app(X1,X2) (0:6) [cnf transformation 58]
58. ! [X0,X1,X2] : var(X0) != app(X1,X2) [rectify 10]
10. ! [X28,X29,X30] : var(X28) != app(X29,X30) [input]
4928. sK6 = X49 {4} (0:3) [sat splitting component]
4927. var(X50) != sK6 {2} (0:4) [sat splitting component]
4926. sK1 = sK12(X47,X48,sK6) | ~tcheck(X47,sK6,X48) {0} (0:10) [sat splitting component]
4928. sK6 = X49 {4} (0:3) [sat splitting component]
4914. sK6 = X432 | sK1 = sK12(X430,X431,sK6) | var(X433) != sK6 | ~tcheck(X430,sK6,X431) (3:17) [superposition 173,995]
995. subst(sK12(X29,X30,X31),X32,X31) = X32 | var(X33) != X31 | ~tcheck(X29,X31,X30) (2:17) [superposition 154,675]
675. var(sK12(X20,X21,X22)) = X22 | var(X23) != X22 | ~tcheck(X20,X22,X21) (1:15) [subsumption resolution 674,326]
326. ~sP0(X16,X15,X14) | var(X17) != X14 (1:8) [superposition 129,118]
118. abs(sK8(X2,X1,X0),sK10(X2,X1,X0),sK9(X2,X1,X0)) = X2 | ~sP0(X0,X1,X2) (0:19) [cnf transformation 94]
94. ! [X0,X1,X2] : (~sP0(X0,X1,X2) | (abs(sK8(X2,X1,X0),sK10(X2,X1,X0),sK9(X2,X1,X0)) = X2 & arrow(sK10(X2,X1,X0),sK11(X2,X1,X0)) = X1 & tcheck(bind(sK8(X2,X1,X0),sK10(X2,X1,X0),X0),sK9(X2,X1,X0),sK11(X2,X1,X0)))) [skolemisation 93]
93. ! [X0,X1,X2] : (~sP0(X0,X1,X2) | ? [X3,X4,X5,X6] : (abs(X3,X5,X4) = X2 & arrow(X5,X6) = X1 & tcheck(bind(X3,X5,X0),X4,X6))) [rectify 92]
92. ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9))) [nnf transformation 87]
87. ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9))) [predicate definition introduction]
129. var(X0) != abs(X1,X2,X3) (0:7) [cnf transformation 57]
57. ! [X0,X1,X2,X3] : var(X0) != abs(X1,X2,X3) [rectify 9]
9. ! [X24,X25,X26,X27] : var(X24) != abs(X25,X26,X27) [input]
674. var(X23) != X22 | sP0(X20,X21,X22) | var(sK12(X20,X21,X22)) = X22 | ~tcheck(X20,X22,X21) (1:19) [superposition 130,121]
121. app(sK13(X2,X1,X0),sK14(X2,X1,X0)) = X0 | sP0(X2,X1,X0) | var(sK12(X2,X1,X0)) = X0 | ~tcheck(X2,X0,X1) (0:26) [cnf transformation 96]
96. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | (var(sK12(X2,X1,X0)) = X0 & some(X1) = lookup(sK12(X2,X1,X0),X2)) | sP0(X2,X1,X0) | (app(sK13(X2,X1,X0),sK14(X2,X1,X0)) = X0 & tcheck(X2,sK13(X2,X1,X0),arrow(sK15(X2,X1,X0),X1)) & tcheck(X2,sK14(X2,X1,X0),sK15(X2,X1,X0)))) [skolemisation 95]
95. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X6] : (var(X6) = X0 & some(X1) = lookup(X6,X2)) | sP0(X2,X1,X0) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))) [rectify 88]
88. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | sP0(X2,X1,X0) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))) [definition folding 83,87]
87. ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9))) [predicate definition introduction]
83. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))) [flattening 82]
82. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | (? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))) [ennf transformation 54]
54. ! [X0,X1,X2] : (tcheck(X2,X0,X1) => (? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))) [rectify 20]
20. ! [X51,X50,X44] : (tcheck(X44,X51,X50) => (? [X40] : (var(X40) = X51 & some(X50) = lookup(X40,X44)) | ? [X40,X54,X55,X56] : (abs(X40,X55,X54) = X51 & arrow(X55,X56) = X50 & tcheck(bind(X40,X55,X44),X54,X56)) | ? [X53,X54,X52] : (app(X53,X54) = X51 & tcheck(X44,X53,arrow(X52,X50)) & tcheck(X44,X54,X52)))) [input]
130. var(X0) != app(X1,X2) (0:6) [cnf transformation 58]
58. ! [X0,X1,X2] : var(X0) != app(X1,X2) [rectify 10]
10. ! [X28,X29,X30] : var(X28) != app(X29,X30) [input]
154. subst(X4,X6,var(X4)) = X6 (0:7) [equality resolution 153]
153. subst(X0,X6,var(X4)) = X6 | X0 != X4 (0:10) [equality resolution 152]
152. subst(X0,X1,var(X4)) = X6 | X1 != X6 | X0 != X4 (0:13) [equality resolution 151]
151. subst(X0,X1,X2) = X6 | var(X4) != X2 | X1 != X6 | X0 != X4 (0:16) [equality resolution 150]
150. X5 = X6 | subst(X0,X1,X2) != X5 | var(X4) != X2 | X1 != X6 | X0 != X4 (0:19) [equality resolution 107]
107. X5 = X6 | X3 != X4 | subst(X0,X1,X2) != X5 | var(X4) != X2 | X1 != X6 | X0 != X3 (0:22) [cnf transformation 70]
70. ! [X0,X1,X2,X3,X4,X5,X6] : (X0 != X3 | X1 != X6 | var(X4) != X2 | subst(X0,X1,X2) != X5 | X3 != X4 | X5 = X6) [flattening 69]
69. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 != X3 | X1 != X6 | var(X4) != X2 | subst(X0,X1,X2) != X5) | (X3 != X4 | X5 = X6)) [ennf transformation 43]
43. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 = X3 & X1 = X6 & var(X4) = X2 & subst(X0,X1,X2) = X5) => (X3 = X4 => X5 = X6)) [true and false elimination 42]
42. ! [X0,X1,X2,X3,X4,X5,X6] : ((X0 = X3 & X1 = X6 & var(X4) = X2 & subst(X0,X1,X2) = X5) => ((X3 = X4 & $true) => X5 = X6)) [rectify 24]
24. ! [X58,X59,X60,X40,X47,X61,X51] : ((X40 = X58 & X51 = X59 & var(X47) = X60 & subst(X58,X59,X60) = X61) => ((X40 = X47 & $true) => X51 = X61)) [input]
173. subst(X0,X1,sK6) = sK6 | sK1 = X0 (1:9) [superposition 149,98]
98. var(sK1) = sK6 (0:4) [cnf transformation 89]
89. var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) [skolemisation 60]
60. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 59]
59. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 31]
31. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
149. var(X6) = subst(X4,X2,var(X6)) | X4 = X6 (0:11) [equality resolution 148]
148. var(X6) = subst(X1,X2,var(X6)) | X4 = X6 | X1 != X4 (0:14) [equality resolution 147]
147. var(X6) = subst(X1,X2,var(X6)) | X4 = X6 | X0 != X2 | X1 != X4 (0:17) [equality resolution 146]
146. var(X6) = subst(X1,X2,X3) | X4 = X6 | var(X6) != X3 | X0 != X2 | X1 != X4 (0:20) [equality resolution 106]
106. var(X6) = X5 | X4 = X6 | subst(X1,X2,X3) != X5 | var(X6) != X3 | X0 != X2 | X1 != X4 (0:23) [cnf transformation 68]
68. ! [X0,X1,X2,X3,X4,X5,X6] : (X1 != X4 | X0 != X2 | var(X6) != X3 | subst(X1,X2,X3) != X5 | X4 = X6 | var(X6) = X5) [flattening 67]
67. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 != X4 | X0 != X2 | var(X6) != X3 | subst(X1,X2,X3) != X5) | (X4 = X6 | var(X6) = X5)) [ennf transformation 41]
41. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X4 & X0 = X2 & var(X6) = X3 & subst(X1,X2,X3) = X5) => (X4 != X6 => var(X6) = X5)) [flattening 40]
40. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X4 & X0 = X2 & var(X6) = X3 & subst(X1,X2,X3) = X5) => (~X4 = X6 => var(X6) = X5)) [true and false elimination 39]
39. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X4 & X0 = X2 & var(X6) = X3 & subst(X1,X2,X3) = X5) => ((~X4 = X6 & $true) => var(X6) = X5)) [rectify 25]
25. ! [X51,X58,X59,X60,X40,X61,X47] : ((X40 = X58 & X51 = X59 & var(X47) = X60 & subst(X58,X59,X60) = X61) => ((~X40 = X47 & $true) => var(X47) = X61)) [input]
10338. some(X1) = lookup(sK1,X0) | ~tcheck(X0,sK6,X1) {46} (0:10) [sat splitting component]
10337. some(X1) = lookup(sK1,X0) | var(X2) != sK6 | ~tcheck(X0,sK6,X1) {0} (2:14) [duplicate literal removal 10333]
10333. some(X1) = lookup(sK1,X0) | var(X2) != sK6 | ~tcheck(X0,sK6,X1) | ~tcheck(X0,sK6,X1) {0} (2:18) [superposition 989,4926]
4926. sK1 = sK12(X47,X48,sK6) | ~tcheck(X47,sK6,X48) {0} (0:10) [sat splitting component]
989. some(X24) = lookup(sK12(X23,X24,X25),X23) | var(X26) != X25 | ~tcheck(X23,X25,X24) (1:17) [subsumption resolution 988,326]
326. ~sP0(X16,X15,X14) | var(X17) != X14 (1:8) [superposition 129,118]
118. abs(sK8(X2,X1,X0),sK10(X2,X1,X0),sK9(X2,X1,X0)) = X2 | ~sP0(X0,X1,X2) (0:19) [cnf transformation 94]
94. ! [X0,X1,X2] : (~sP0(X0,X1,X2) | (abs(sK8(X2,X1,X0),sK10(X2,X1,X0),sK9(X2,X1,X0)) = X2 & arrow(sK10(X2,X1,X0),sK11(X2,X1,X0)) = X1 & tcheck(bind(sK8(X2,X1,X0),sK10(X2,X1,X0),X0),sK9(X2,X1,X0),sK11(X2,X1,X0)))) [skolemisation 93]
93. ! [X0,X1,X2] : (~sP0(X0,X1,X2) | ? [X3,X4,X5,X6] : (abs(X3,X5,X4) = X2 & arrow(X5,X6) = X1 & tcheck(bind(X3,X5,X0),X4,X6))) [rectify 92]
92. ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9))) [nnf transformation 87]
87. ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9))) [predicate definition introduction]
129. var(X0) != abs(X1,X2,X3) (0:7) [cnf transformation 57]
57. ! [X0,X1,X2,X3] : var(X0) != abs(X1,X2,X3) [rectify 9]
9. ! [X24,X25,X26,X27] : var(X24) != abs(X25,X26,X27) [input]
988. var(X26) != X25 | sP0(X23,X24,X25) | some(X24) = lookup(sK12(X23,X24,X25),X23) | ~tcheck(X23,X25,X24) (1:21) [superposition 130,122]
122. app(sK13(X2,X1,X0),sK14(X2,X1,X0)) = X0 | sP0(X2,X1,X0) | some(X1) = lookup(sK12(X2,X1,X0),X2) | ~tcheck(X2,X0,X1) (0:28) [cnf transformation 96]
96. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | (var(sK12(X2,X1,X0)) = X0 & some(X1) = lookup(sK12(X2,X1,X0),X2)) | sP0(X2,X1,X0) | (app(sK13(X2,X1,X0),sK14(X2,X1,X0)) = X0 & tcheck(X2,sK13(X2,X1,X0),arrow(sK15(X2,X1,X0),X1)) & tcheck(X2,sK14(X2,X1,X0),sK15(X2,X1,X0)))) [skolemisation 95]
95. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X6] : (var(X6) = X0 & some(X1) = lookup(X6,X2)) | sP0(X2,X1,X0) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))) [rectify 88]
88. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | sP0(X2,X1,X0) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))) [definition folding 83,87]
87. ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9))) [predicate definition introduction]
83. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))) [flattening 82]
82. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | (? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))) [ennf transformation 54]
54. ! [X0,X1,X2] : (tcheck(X2,X0,X1) => (? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))) [rectify 20]
20. ! [X51,X50,X44] : (tcheck(X44,X51,X50) => (? [X40] : (var(X40) = X51 & some(X50) = lookup(X40,X44)) | ? [X40,X54,X55,X56] : (abs(X40,X55,X54) = X51 & arrow(X55,X56) = X50 & tcheck(bind(X40,X55,X44),X54,X56)) | ? [X53,X54,X52] : (app(X53,X54) = X51 & tcheck(X44,X53,arrow(X52,X50)) & tcheck(X44,X54,X52)))) [input]
130. var(X0) != app(X1,X2) (0:6) [cnf transformation 58]
58. ! [X0,X1,X2] : var(X0) != app(X1,X2) [rectify 10]
10. ! [X28,X29,X30] : var(X28) != app(X29,X30) [input]
10444. $false {46} (4:0) [subsumption resolution 10441,180]
180. ~tcheck(sK3,sK6,sK7) (2:4) [subsumption resolution 178,99]
99. sK1 != sK4 (0:3) [cnf transformation 89]
89. var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) [skolemisation 60]
60. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 59]
59. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 31]
31. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
178. ~tcheck(sK3,sK6,sK7) | sK1 = sK4 (2:7) [superposition 102,173]
173. subst(X0,X1,sK6) = sK6 | sK1 = X0 (1:9) [superposition 149,98]
98. var(sK1) = sK6 (0:4) [cnf transformation 89]
89. var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) [skolemisation 60]
60. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 59]
59. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 31]
31. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
149. var(X6) = subst(X4,X2,var(X6)) | X4 = X6 (0:11) [equality resolution 148]
148. var(X6) = subst(X1,X2,var(X6)) | X4 = X6 | X1 != X4 (0:14) [equality resolution 147]
147. var(X6) = subst(X1,X2,var(X6)) | X4 = X6 | X0 != X2 | X1 != X4 (0:17) [equality resolution 146]
146. var(X6) = subst(X1,X2,X3) | X4 = X6 | var(X6) != X3 | X0 != X2 | X1 != X4 (0:20) [equality resolution 106]
106. var(X6) = X5 | X4 = X6 | subst(X1,X2,X3) != X5 | var(X6) != X3 | X0 != X2 | X1 != X4 (0:23) [cnf transformation 68]
68. ! [X0,X1,X2,X3,X4,X5,X6] : (X1 != X4 | X0 != X2 | var(X6) != X3 | subst(X1,X2,X3) != X5 | X4 = X6 | var(X6) = X5) [flattening 67]
67. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 != X4 | X0 != X2 | var(X6) != X3 | subst(X1,X2,X3) != X5) | (X4 = X6 | var(X6) = X5)) [ennf transformation 41]
41. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X4 & X0 = X2 & var(X6) = X3 & subst(X1,X2,X3) = X5) => (X4 != X6 => var(X6) = X5)) [flattening 40]
40. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X4 & X0 = X2 & var(X6) = X3 & subst(X1,X2,X3) = X5) => (~X4 = X6 => var(X6) = X5)) [true and false elimination 39]
39. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X4 & X0 = X2 & var(X6) = X3 & subst(X1,X2,X3) = X5) => ((~X4 = X6 & $true) => var(X6) = X5)) [rectify 25]
25. ! [X51,X58,X59,X60,X40,X61,X47] : ((X40 = X58 & X51 = X59 & var(X47) = X60 & subst(X58,X59,X60) = X61) => ((~X40 = X47 & $true) => var(X47) = X61)) [input]
102. ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) (0:7) [cnf transformation 89]
89. var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) [skolemisation 60]
60. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 59]
59. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 31]
31. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
10441. tcheck(sK3,sK6,sK7) {46} (4:4) [equality resolution 10414]
10414. some(sK7) != some(X5) | tcheck(sK3,sK6,X5) {46} (3:9) [forward demodulation 10413,98]
98. var(sK1) = sK6 (0:4) [cnf transformation 89]
89. var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) [skolemisation 60]
60. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 59]
59. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 31]
31. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
10413. some(sK7) != some(X5) | tcheck(sK3,var(sK1),X5) {46} (3:10) [superposition 128,10406]
10406. some(sK7) = lookup(sK1,sK3) {46} (2:6) [subsumption resolution 10402,99]
99. sK1 != sK4 (0:3) [cnf transformation 89]
89. var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) [skolemisation 60]
60. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 59]
59. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 31]
31. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
10402. some(sK7) = lookup(sK1,sK3) | sK1 = sK4 {46} (2:9) [resolution 10341,101]
101. tcheck(bind(sK4,sK2,sK3),sK6,sK7) (0:7) [cnf transformation 89]
89. var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7) [skolemisation 60]
60. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 59]
59. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 31]
31. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
10341. ~tcheck(bind(X3,X4,X5),sK6,X6) | some(X6) = lookup(sK1,X5) | sK1 = X3 {46} (1:16) [superposition 10338,158]
158. lookup(X5,X6) = lookup(X5,bind(X3,X0,X6)) | X3 = X5 (0:13) [equality resolution 157]
157. lookup(X5,X6) = lookup(X1,bind(X3,X0,X6)) | X3 = X5 | X1 != X5 (0:16) [equality resolution 156]
156. lookup(X1,X2) = lookup(X5,X6) | X3 = X5 | bind(X3,X0,X6) != X2 | X1 != X5 (0:19) [equality resolution 109]
109. lookup(X5,X6) = X4 | X3 = X5 | lookup(X1,X2) != X4 | bind(X3,X0,X6) != X2 | X1 != X5 (0:22) [cnf transformation 74]
74. ! [X0,X1,X2,X3,X4,X5,X6] : (X1 != X5 | bind(X3,X0,X6) != X2 | lookup(X1,X2) != X4 | X3 = X5 | lookup(X5,X6) = X4) [flattening 73]
73. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 != X5 | bind(X3,X0,X6) != X2 | lookup(X1,X2) != X4) | (X3 = X5 | lookup(X5,X6) = X4)) [ennf transformation 47]
47. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X5 & bind(X3,X0,X6) = X2 & lookup(X1,X2) = X4) => (X3 != X5 => lookup(X5,X6) = X4)) [flattening 46]
46. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X5 & bind(X3,X0,X6) = X2 & lookup(X1,X2) = X4) => (~X3 = X5 => lookup(X5,X6) = X4)) [true and false elimination 45]
45. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X5 & bind(X3,X0,X6) = X2 & lookup(X1,X2) = X4) => ((~X3 = X5 & $true) => lookup(X5,X6) = X4)) [rectify 16]
16. ! [X49,X45,X46,X47,X48,X40,X44] : ((X40 = X45 & bind(X47,X49,X44) = X46 & lookup(X45,X46) = X48) => ((~X40 = X47 & $true) => lookup(X40,X44) = X48)) [input]
10338. some(X1) = lookup(sK1,X0) | ~tcheck(X0,sK6,X1) {46} (0:10) [sat splitting component]
128. some(X2) != lookup(X1,X0) | tcheck(X0,var(X1),X2) (0:11) [cnf transformation 86]
86. ! [X0,X1,X2] : (some(X2) != lookup(X1,X0) | tcheck(X0,var(X1),X2)) [ennf transformation 56]
56. ! [X0,X1,X2] : (some(X2) = lookup(X1,X0) => tcheck(X0,var(X1),X2)) [rectify 17]
17. ! [X44,X40,X50] : (some(X50) = lookup(X40,X44) => tcheck(X44,var(X40),X50)) [input]
% SZS output end Proof for Subst-1
% ------------------------------
% Version: Vampire 3.0 (commit e3e1a68)
% Termination reason: Refutation

% Active clauses: 283
% Passive clauses: 3523
% Generated clauses: 10065
% Final active clauses: 97
% Final passive clauses: 2715
% Input formulas: 29
% Initial clauses: 35
% 
% Selected by SInE selection: 19
% SInE iterations: 1
% 
% Duplicate literals: 59
% Trivial inequalities: 3
% Fw subsumption resolutions: 52
% Fw demodulations: 1878
% Bw demodulations: 177
% 
% Simple tautologies: 11
% Equational tautologies: 14
% Forward subsumptions: 4814
% Fw demodulations to eq. taut.: 181
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 116
% Forward superposition: 4264
% Backward superposition: 3194
% Self superposition: 120
% Equality resolution: 16
% 
% Split clauses: 20
% Split components: 52
% SAT solver clauses: 50
% SAT solver unit clauses: 16
% SAT solver binary clauses: 7
% SAT solver learnt clauses: 1
% 
% Sat splits: 20
% Sat splitting refutations: 16
% 
% Memory used [KB]: 6268
% Time elapsed: 0.144 s
% ------------------------------
% ------------------------------
% Success in time 0.549 s
