% START OF SYSTEM OUTPUT
Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 599 next slice time: 3
lrs-1010_7_bs=unit_only:bsr=unit_only:cond=fast:drc=off:ep=on:flr=on:lcm=predicate:nwc=1.7:stl=30:sgo=on:spo=on:spl=sat:ssfq=1.1:ssnc=none:sp=occurrence_1 on SOT_HJC5
% (3775)Time limit reached!
% ------------------------------
% Version: Vampire 3.0 (revision 1803, CASC-24-only license, for any other Vampire license check http://vprover.org/)
% SZS status Timeout
% Termination phase: Saturation

% Active clauses: 302
% Passive clauses: 2560
% Generated clauses: 10872
% Final active clauses: 299
% Final passive clauses: 2225
% Input formulas: 29
% Initial clauses: 55
% 
% Duplicate literals: 422
% Fw subsumption resolutions: 100
% Fw demodulations: 889
% Bw demodulations: 6
% Condensations: 278
% 
% Simple tautologies: 1
% Equational tautologies: 9
% Forward subsumptions: 6386
% Backward subsumptions: 30
% Fw demodulations to eq. taut.: 111
% 
% Binary resolution: 445
% Factoring: 2
% Forward superposition: 4336
% Backward superposition: 4329
% Self superposition: 48
% 
% Memory used [KB]: 6140
% Time elapsed: 0.400 s
% ------------------------------
% ------------------------------
% remaining time: 595 next slice time: 3
dis-1002_6_bs=off:drc=off:nwc=1.1:nicw=on:sd=1:ss=axioms:st=3.0:sio=off:spl=sat:sser=off_1 on SOT_HJC5
% (3777)Time limit reached!
% ------------------------------
% Version: Vampire 3.0 (revision 1803, CASC-24-only license, for any other Vampire license check http://vprover.org/)
% SZS status Timeout
% Termination phase: Saturation

% Active clauses: 237
% Passive clauses: 5744
% Generated clauses: 11431
% Final active clauses: 233
% Final passive clauses: 5273
% Input formulas: 29
% Initial clauses: 35
% 
% Selected by SInE selection: 19
% SInE iterations: 1
% 
% Duplicate literals: 167
% Fw subsumption resolutions: 4
% Fw demodulations: 2700
% Bw demodulations: 4
% 
% Simple tautologies: 6
% Equational tautologies: 9
% Forward subsumptions: 2619
% Fw demodulations to eq. taut.: 212
% 
% Binary resolution: 137
% Forward superposition: 3846
% Backward superposition: 4473
% Self superposition: 65
% Equality resolution: 12
% 
% Memory used [KB]: 6268
% Time elapsed: 0.400 s
% ------------------------------
% ------------------------------
% remaining time: 591 next slice time: 6
ott+1_5:1_bs=off:bsr=unit_only:cond=fast:drc=off:nwc=5:ss=axioms:sos=all_4 on SOT_HJC5
% (3778)Time limit reached!
% ------------------------------
% Version: Vampire 3.0 (revision 1803, CASC-24-only license, for any other Vampire license check http://vprover.org/)
% SZS status Timeout
% Termination phase: Saturation

% Active clauses: 170
% Passive clauses: 3932
% Generated clauses: 9622
% Final active clauses: 170
% Final passive clauses: 3771
% Input formulas: 29
% Initial clauses: 28
% 
% Selected by SInE selection: 15
% SInE iterations: 3
% 
% Duplicate literals: 467
% Fw subsumption resolutions: 14
% Fw demodulations: 972
% Condensations: 19
% 
% Equational tautologies: 6
% Forward subsumptions: 4088
% Fw demodulations to eq. taut.: 128
% Bw demodulations to eq. taut.: 10
% 
% Forward superposition: 4243
% Backward superposition: 3909
% Self superposition: 48
% Equality resolution: 12
% 
% Memory used [KB]: 9850
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 584 next slice time: 28
ott+2_20_cond=fast:drc=off:flr=on:lcm=reverse:nwc=1.1:nicw=on:sio=off:spl=sat:sser=off:ssfp=10000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_25 on SOT_HJC5
% (3779)Time limit reached!
% ------------------------------
% Version: Vampire 3.0 (revision 1803, CASC-24-only license, for any other Vampire license check http://vprover.org/)
% SZS status Timeout
% Termination phase: Saturation

% Active clauses: 999
% Passive clauses: 11819
% Generated clauses: 54305
% Final active clauses: 963
% Final passive clauses: 7836
% Input formulas: 29
% Initial clauses: 53
% 
% Duplicate literals: 1540
% Fw subsumption resolutions: 154
% Fw demodulations: 5176
% Bw demodulations: 27
% Condensations: 5645
% 
% Simple tautologies: 6
% Equational tautologies: 38
% Forward subsumptions: 29351
% Backward subsumptions: 2993
% Fw demodulations to eq. taut.: 58
% 
% Binary resolution: 219
% Forward superposition: 29192
% Backward superposition: 12149
% Self superposition: 68
% Equality resolution: 88
% 
% Memory used [KB]: 26865
% Time elapsed: 2.900 s
% ------------------------------
% ------------------------------
% remaining time: 555 next slice time: 3
lrs+1011_8:1_bd=off:bs=off:bsr=unit_only:cond=fast:drc=off:fde=none:nwc=1:nicw=on:stl=240:sos=all:sio=off:spo=on:ssac=none:ssfp=1000:ssfq=1.1:ssnc=all_dependent_1 on SOT_HJC5
% (3780)Time limit reached!
% ------------------------------
% Version: Vampire 3.0 (revision 1803, CASC-24-only license, for any other Vampire license check http://vprover.org/)
% SZS status Timeout
% Termination phase: Saturation

% Active clauses: 377
% Passive clauses: 3777
% Generated clauses: 5491
% Final active clauses: 377
% Final passive clauses: 3440
% Input formulas: 29
% Initial clauses: 53
% 
% Duplicate literals: 137
% Trivial inequalities: 5
% Fw subsumption resolutions: 161
% Fw demodulations: 141
% 
% Simple tautologies: 21
% Equational tautologies: 77
% Forward subsumptions: 1119
% Fw demodulations to eq. taut.: 26
% 
% Binary resolution: 3472
% Forward superposition: 931
% Backward superposition: 581
% Self superposition: 8
% 
% Split clauses: 45
% Split components: 111
% Introduced splitting names: 19
% 
% Memory used [KB]: 6012
% Time elapsed: 0.400 s
% ------------------------------
% ------------------------------
% remaining time: 551 next slice time: 107
dis+1_2:3_bs=off:drc=off:lcm=predicate:nwc=3:nicw=on:ss=included:st=1.5:sos=on:sagn=off:sio=off:spl=sat:sser=off:sp=occurrence_97 on SOT_HJC5
% 10 wall clock seconds gone, waiting for system output
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_HJC5
% SZS output start Proof for SOT_HJC5
fof(f165475,plain,(
  $false),
  inference(subsumption_resolution,[],[f165474,f242])).
fof(f242,plain,(
  ~tcheck(sK3,sK6,sK7)),
  inference(subsumption_resolution,[],[f241,f116])).
fof(f116,plain,(
  sK1 != sK4),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  var(sK1) = sK6 & sK1 != sK4 & tcheck(sK3,sK5,sK2) & tcheck(bind(sK4,sK2,sK3),sK6,sK7) & ~tcheck(sK3,subst(sK4,sK5,sK6),sK7)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4,sK5,sK6,sK7])],[f68])).
fof(f68,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6))),
  inference(rectify,[],[f30])).
fof(f30,negated_conjecture,(
  ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56))),
  inference(negated_conjecture,[],[f29])).
fof(f29,conjecture,(
  ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56))),
  file('/tmp/SystemOnTPTPFormReply3740/SOT_HJC5ks','T-subst-var-2')).
fof(f241,plain,(
  ~tcheck(sK3,sK6,sK7) | sK1 = sK4),
  inference(superposition,[],[f119,f214])).
fof(f214,plain,(
  ( ! [X2,X1] : (subst(X1,X2,sK6) = sK6 | sK1 = X1) )),
  inference(superposition,[],[f186,f115])).
fof(f115,plain,(
  var(sK1) = sK6),
  inference(cnf_transformation,[],[f99])).
fof(f186,plain,(
  ( ! [X6,X4,X2] : (var(X6) = subst(X4,X2,var(X6)) | X4 = X6) )),
  inference(equality_resolution,[],[f185])).
fof(f185,plain,(
  ( ! [X6,X4,X2,X1] : (var(X6) = subst(X1,X2,var(X6)) | X4 = X6 | X1 != X4) )),
  inference(equality_resolution,[],[f184])).
fof(f184,plain,(
  ( ! [X6,X4,X2,X0,X1] : (var(X6) = subst(X1,X2,var(X6)) | X4 = X6 | X0 != X2 | X1 != X4) )),
  inference(equality_resolution,[],[f183])).
fof(f183,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (var(X6) = subst(X1,X2,X3) | X4 = X6 | var(X6) != X3 | X0 != X2 | X1 != X4) )),
  inference(equality_resolution,[],[f160])).
fof(f160,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (var(X6) = X5 | X4 = X6 | subst(X1,X2,X3) != X5 | var(X6) != X3 | X0 != X2 | X1 != X4) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (X1 != X4 | X0 != X2 | var(X6) != X3 | subst(X1,X2,X3) != X5 | X4 = X6 | var(X6) = X5)),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 != X4 | X0 != X2 | var(X6) != X3 | subst(X1,X2,X3) != X5) | (X4 = X6 | var(X6) = X5))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X4 & X0 = X2 & var(X6) = X3 & subst(X1,X2,X3) = X5) => (X4 != X6 => var(X6) = X5))),
  inference(true_and_false_elimination,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X4 & X0 = X2 & var(X6) = X3 & subst(X1,X2,X3) = X5) => ((X4 != X6 & $true) => var(X6) = X5))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X51,X58,X59,X60,X40,X61,X47] : ((X40 = X58 & X51 = X59 & var(X47) = X60 & subst(X58,X59,X60) = X61) => ((X40 != X47 & $true) => var(X47) = X61))),
  file('/tmp/SystemOnTPTPFormReply3740/SOT_HJC5ks',subst29)).
fof(f119,plain,(
  ~tcheck(sK3,subst(sK4,sK5,sK6),sK7)),
  inference(cnf_transformation,[],[f99])).
fof(f165474,plain,(
  tcheck(sK3,sK6,sK7)),
  inference(equality_resolution,[],[f165473])).
fof(f165473,plain,(
  ( ! [X12] : (some(sK7) != some(X12) | tcheck(sK3,sK6,X12)) )),
  inference(forward_demodulation,[],[f165472,f115])).
fof(f165472,plain,(
  ( ! [X12] : (some(sK7) != some(X12) | tcheck(sK3,var(sK1),X12)) )),
  inference(superposition,[],[f138,f165457])).
fof(f165457,plain,(
  some(sK7) = lookup(sK1,sK3)),
  inference(subsumption_resolution,[],[f165453,f116])).
fof(f165453,plain,(
  some(sK7) = lookup(sK1,sK3) | sK1 = sK4),
  inference(superposition,[],[f165415,f194])).
fof(f194,plain,(
  ( ! [X6,X0,X5,X3] : (lookup(X5,X6) = lookup(X5,bind(X3,X0,X6)) | X3 = X5) )),
  inference(equality_resolution,[],[f193])).
fof(f193,plain,(
  ( ! [X6,X0,X5,X3,X1] : (lookup(X5,X6) = lookup(X1,bind(X3,X0,X6)) | X3 = X5 | X1 != X5) )),
  inference(equality_resolution,[],[f192])).
fof(f192,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (lookup(X1,X2) = lookup(X5,X6) | X3 = X5 | bind(X3,X0,X6) != X2 | X1 != X5) )),
  inference(equality_resolution,[],[f162])).
fof(f162,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (lookup(X5,X6) = X4 | X3 = X5 | lookup(X1,X2) != X4 | bind(X3,X0,X6) != X2 | X1 != X5) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (X1 != X5 | bind(X3,X0,X6) != X2 | lookup(X1,X2) != X4 | X3 = X5 | lookup(X5,X6) = X4)),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 != X5 | bind(X3,X0,X6) != X2 | lookup(X1,X2) != X4) | (X3 = X5 | lookup(X5,X6) = X4))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X5 & bind(X3,X0,X6) = X2 & lookup(X1,X2) = X4) => (X3 != X5 => lookup(X5,X6) = X4))),
  inference(true_and_false_elimination,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X5 & bind(X3,X0,X6) = X2 & lookup(X1,X2) = X4) => ((X3 != X5 & $true) => lookup(X5,X6) = X4))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X49,X45,X46,X47,X48,X40,X44] : ((X40 = X45 & bind(X47,X49,X44) = X46 & lookup(X45,X46) = X48) => ((X40 != X47 & $true) => lookup(X40,X44) = X48))),
  file('/tmp/SystemOnTPTPFormReply3740/SOT_HJC5ks',lookup16)).
fof(f165415,plain,(
  some(sK7) = lookup(sK1,bind(sK4,sK2,sK3))),
  inference(subsumption_resolution,[],[f165414,f160664])).
fof(f160664,plain,(
  ( ! [X2,X3] : (app(X2,X3) != sK6) )),
  inference(superposition,[],[f127,f160659])).
fof(f160659,plain,(
  var(sK12(bind(sK4,sK2,sK3),sK7,sK6)) = sK6),
  inference(trivial_inequality_removal,[],[f160658])).
fof(f160658,plain,(
  sK6 != sK6 | var(sK12(bind(sK4,sK2,sK3),sK7,sK6)) = sK6),
  inference(superposition,[],[f160557,f115])).
fof(f160557,plain,(
  ( ! [X0] : (var(X0) != sK6 | var(sK12(bind(sK4,sK2,sK3),sK7,sK6)) = sK6) )),
  inference(superposition,[],[f127,f3816])).
fof(f3816,plain,(
  app(sK13(bind(sK4,sK2,sK3),sK7,sK6),sK14(bind(sK4,sK2,sK3),sK7,sK6)) = sK6 | var(sK12(bind(sK4,sK2,sK3),sK7,sK6)) = sK6),
  inference(trivial_inequality_removal,[],[f3815])).
fof(f3815,plain,(
  sK6 != sK6 | app(sK13(bind(sK4,sK2,sK3),sK7,sK6),sK14(bind(sK4,sK2,sK3),sK7,sK6)) = sK6 | var(sK12(bind(sK4,sK2,sK3),sK7,sK6)) = sK6),
  inference(superposition,[],[f1246,f115])).
fof(f1246,plain,(
  ( ! [X0] : (var(X0) != sK6 | app(sK13(bind(sK4,sK2,sK3),sK7,sK6),sK14(bind(sK4,sK2,sK3),sK7,sK6)) = sK6 | var(sK12(bind(sK4,sK2,sK3),sK7,sK6)) = sK6) )),
  inference(superposition,[],[f139,f371])).
fof(f371,plain,(
  abs(sK8(sK6,sK7,bind(sK4,sK2,sK3)),sK10(sK6,sK7,bind(sK4,sK2,sK3)),sK9(sK6,sK7,bind(sK4,sK2,sK3))) = sK6 | app(sK13(bind(sK4,sK2,sK3),sK7,sK6),sK14(bind(sK4,sK2,sK3),sK7,sK6)) = sK6 | var(sK12(bind(sK4,sK2,sK3),sK7,sK6)) = sK6),
  inference(resolution,[],[f226,f129])).
fof(f129,plain,(
  ( ! [X2,X0,X1] : (~sP0(X0,X1,X2) | abs(sK8(X2,X1,X0),sK10(X2,X1,X0),sK9(X2,X1,X0)) = X2) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : (~sP0(X0,X1,X2) | (abs(sK8(X2,X1,X0),sK10(X2,X1,X0),sK9(X2,X1,X0)) = X2 & arrow(sK10(X2,X1,X0),sK11(X2,X1,X0)) = X1 & tcheck(bind(sK8(X2,X1,X0),sK10(X2,X1,X0),X0),sK9(X2,X1,X0),sK11(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10,sK11])],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : (~sP0(X0,X1,X2) | ? [X3,X4,X5,X6] : (abs(X3,X5,X4) = X2 & arrow(X5,X6) = X1 & tcheck(bind(X3,X5,X0),X4,X6)))),
  inference(rectify,[],[f102])).
fof(f102,plain,(
  ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)))),
  inference(nnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f226,plain,(
  sP0(bind(sK4,sK2,sK3),sK7,sK6) | var(sK12(bind(sK4,sK2,sK3),sK7,sK6)) = sK6 | app(sK13(bind(sK4,sK2,sK3),sK7,sK6),sK14(bind(sK4,sK2,sK3),sK7,sK6)) = sK6),
  inference(resolution,[],[f118,f132])).
fof(f132,plain,(
  ( ! [X2,X0,X1] : (~tcheck(X2,X0,X1) | sP0(X2,X1,X0) | var(sK12(X2,X1,X0)) = X0 | app(sK13(X2,X1,X0),sK14(X2,X1,X0)) = X0) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | (var(sK12(X2,X1,X0)) = X0 & some(X1) = lookup(sK12(X2,X1,X0),X2)) | sP0(X2,X1,X0) | (app(sK13(X2,X1,X0),sK14(X2,X1,X0)) = X0 & tcheck(X2,sK13(X2,X1,X0),arrow(sK15(X2,X1,X0),X1)) & tcheck(X2,sK14(X2,X1,X0),sK15(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X6] : (var(X6) = X0 & some(X1) = lookup(X6,X2)) | sP0(X2,X1,X0) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))),
  inference(rectify,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | sP0(X2,X1,X0) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))),
  inference(definition_folding,[],[f70,f97])).
fof(f70,plain,(
  ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | (? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : (tcheck(X2,X0,X1) => (? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X51,X50,X44] : (tcheck(X44,X51,X50) => (? [X40] : (var(X40) = X51 & some(X50) = lookup(X40,X44)) | ? [X40,X54,X55,X56] : (abs(X40,X55,X54) = X51 & arrow(X55,X56) = X50 & tcheck(bind(X40,X55,X44),X54,X56)) | ? [X53,X54,X52] : (app(X53,X54) = X51 & tcheck(X44,X53,arrow(X52,X50)) & tcheck(X44,X54,X52))))),
  file('/tmp/SystemOnTPTPFormReply3740/SOT_HJC5ks','T-inv')).
fof(f118,plain,(
  tcheck(bind(sK4,sK2,sK3),sK6,sK7)),
  inference(cnf_transformation,[],[f99])).
fof(f139,plain,(
  ( ! [X2,X0,X3,X1] : (var(X0) != abs(X1,X2,X3)) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3] : var(X0) != abs(X1,X2,X3)),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X24,X25,X26,X27] : var(X24) != abs(X25,X26,X27)),
  file('/tmp/SystemOnTPTPFormReply3740/SOT_HJC5ks','DIFF-var-abs')).
fof(f127,plain,(
  ( ! [X2,X0,X1] : (var(X0) != app(X1,X2)) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2] : var(X0) != app(X1,X2)),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X28,X29,X30] : var(X28) != app(X29,X30)),
  file('/tmp/SystemOnTPTPFormReply3740/SOT_HJC5ks','DIFF-var-app')).
fof(f165414,plain,(
  some(sK7) = lookup(sK1,bind(sK4,sK2,sK3)) | app(sK13(bind(sK4,sK2,sK3),sK7,sK6),sK14(bind(sK4,sK2,sK3),sK7,sK6)) = sK6),
  inference(subsumption_resolution,[],[f165328,f160665])).
fof(f160665,plain,(
  ( ! [X6,X4,X5] : (abs(X4,X5,X6) != sK6) )),
  inference(superposition,[],[f139,f160659])).
fof(f165328,plain,(
  some(sK7) = lookup(sK1,bind(sK4,sK2,sK3)) | abs(sK8(sK6,sK7,bind(sK4,sK2,sK3)),sK10(sK6,sK7,bind(sK4,sK2,sK3)),sK9(sK6,sK7,bind(sK4,sK2,sK3))) = sK6 | app(sK13(bind(sK4,sK2,sK3),sK7,sK6),sK14(bind(sK4,sK2,sK3),sK7,sK6)) = sK6),
  inference(backward_demodulation,[],[f160841,f380])).
fof(f380,plain,(
  abs(sK8(sK6,sK7,bind(sK4,sK2,sK3)),sK10(sK6,sK7,bind(sK4,sK2,sK3)),sK9(sK6,sK7,bind(sK4,sK2,sK3))) = sK6 | app(sK13(bind(sK4,sK2,sK3),sK7,sK6),sK14(bind(sK4,sK2,sK3),sK7,sK6)) = sK6 | some(sK7) = lookup(sK12(bind(sK4,sK2,sK3),sK7,sK6),bind(sK4,sK2,sK3))),
  inference(resolution,[],[f227,f129])).
fof(f227,plain,(
  sP0(bind(sK4,sK2,sK3),sK7,sK6) | some(sK7) = lookup(sK12(bind(sK4,sK2,sK3),sK7,sK6),bind(sK4,sK2,sK3)) | app(sK13(bind(sK4,sK2,sK3),sK7,sK6),sK14(bind(sK4,sK2,sK3),sK7,sK6)) = sK6),
  inference(resolution,[],[f118,f133])).
fof(f133,plain,(
  ( ! [X2,X0,X1] : (~tcheck(X2,X0,X1) | sP0(X2,X1,X0) | some(X1) = lookup(sK12(X2,X1,X0),X2) | app(sK13(X2,X1,X0),sK14(X2,X1,X0)) = X0) )),
  inference(cnf_transformation,[],[f106])).
fof(f160841,plain,(
  sK1 = sK12(bind(sK4,sK2,sK3),sK7,sK6)),
  inference(trivial_inequality_removal,[],[f160668])).
fof(f160668,plain,(
  sK6 != sK6 | sK1 = sK12(bind(sK4,sK2,sK3),sK7,sK6)),
  inference(superposition,[],[f215,f160659])).
fof(f215,plain,(
  ( ! [X8] : (var(X8) != sK6 | sK1 = X8) )),
  inference(superposition,[],[f123,f115])).
fof(f123,plain,(
  ( ! [X0,X1] : (var(X0) != var(X1) | X0 = X1) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ((var(X0) != var(X1) | X0 = X1) & (X0 != X1 | var(X0) = var(X1)))),
  inference(nnf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1] : (var(X0) = var(X1) <=> X0 = X1)),
  file('/tmp/SystemOnTPTPFormReply3740/SOT_HJC5ks','EQ-var')).
fof(f138,plain,(
  ( ! [X2,X0,X1] : (some(X2) != lookup(X1,X0) | tcheck(X0,var(X1),X2)) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : (some(X2) != lookup(X1,X0) | tcheck(X0,var(X1),X2))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2] : (some(X2) = lookup(X1,X0) => tcheck(X0,var(X1),X2))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X44,X40,X50] : (some(X50) = lookup(X40,X44) => tcheck(X44,var(X40),X50))),
  file('/tmp/SystemOnTPTPFormReply3740/SOT_HJC5ks','T-var')).
% SZS output end Proof for SOT_HJC5
% ------------------------------
% Version: Vampire 3.0 (revision 1803, CASC-24-only license, for any other Vampire license check http://vprover.org/)
% SZS status Unsatisfiable

% Active clauses: 1179
% Passive clauses: 87239
% Generated clauses: 139904
% Final active clauses: 980
% Final passive clauses: 85343
% Input formulas: 29
% Initial clauses: 53
% 
% Selected by SInE selection: 29
% SInE iterations: 1
% 
% Duplicate literals: 4724
% Trivial inequalities: 26
% Fw subsumption resolutions: 767
% Fw demodulations: 19187
% Bw demodulations: 106
% 
% Equational tautologies: 13
% Forward subsumptions: 28121
% Fw demodulations to eq. taut.: 615
% Bw demodulations to eq. taut.: 93
% 
% Binary resolution: 1145
% Forward superposition: 59667
% Backward superposition: 54084
% Self superposition: 166
% Equality resolution: 51
% 
% Memory used [KB]: 85840
% Time elapsed: 6.847 s
% ------------------------------
% ------------------------------
% Success in time 11.667 s

% END OF SYSTEM OUTPUT
% RESULT: SOT_HJC5ks - Vampire---3.0 says Theorem - CPU = 11.60 WC = 11.65 
% OUTPUT: SOT_HJC5ks - Vampire---3.0 says Refutation - CPU = 11.60 WC = 11.66 

