Refutation found. Thanks to Tanya!
1155. $false (0:0) [sat splitting refutation 1154,368,384,382,380,1108,432,434,384,458]
458. $false {12} (3:0) [subsumption resolution 457,235]
235. ~tcheck(sK11,sK14,sK15) (2:4) [subsumption resolution 233,167]
167. sK9 != sK12 (0:3) [cnf transformation 117]
117. var(sK9) = sK14 & sK9 != sK12 & tcheck(sK11,sK13,sK10) & tcheck(bind(sK12,sK10,sK11),sK14,sK15) & ~tcheck(sK11,subst(sK12,sK13,sK14),sK15) [skolemisation 99]
99. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 98]
98. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 69]
69. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
233. ~tcheck(sK11,sK14,sK15) | sK9 = sK12 (2:7) [superposition 170,228]
228. subst(X0,X1,sK14) = sK14 | sK9 = X0 (1:9) [superposition 202,166]
166. var(sK9) = sK14 (0:4) [cnf transformation 117]
117. var(sK9) = sK14 & sK9 != sK12 & tcheck(sK11,sK13,sK10) & tcheck(bind(sK12,sK10,sK11),sK14,sK15) & ~tcheck(sK11,subst(sK12,sK13,sK14),sK15) [skolemisation 99]
99. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 98]
98. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 69]
69. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
202. var(X6) = subst(X4,X2,var(X6)) | X4 = X6 (0:11) [equality resolution 201]
201. var(X6) = subst(X1,X2,var(X6)) | X4 = X6 | X1 != X4 (0:14) [equality resolution 200]
200. var(X6) = subst(X1,X2,var(X6)) | X4 = X6 | X0 != X2 | X1 != X4 (0:17) [equality resolution 199]
199. var(X6) = subst(X1,X2,X3) | X4 = X6 | var(X6) != X3 | X0 != X2 | X1 != X4 (0:20) [equality resolution 162]
162. var(X6) = X5 | X4 = X6 | subst(X1,X2,X3) != X5 | var(X6) != X3 | X0 != X2 | X1 != X4 (0:23) [cnf transformation 91]
91. ! [X0,X1,X2,X3,X4,X5,X6] : (X1 != X4 | X0 != X2 | var(X6) != X3 | subst(X1,X2,X3) != X5 | X4 = X6 | var(X6) = X5) [flattening 90]
90. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 != X4 | X0 != X2 | var(X6) != X3 | subst(X1,X2,X3) != X5) | (X4 = X6 | var(X6) = X5)) [ennf transformation 61]
61. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X4 & X0 = X2 & var(X6) = X3 & subst(X1,X2,X3) = X5) => (X4 != X6 => var(X6) = X5)) [flattening 60]
60. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X4 & X0 = X2 & var(X6) = X3 & subst(X1,X2,X3) = X5) => (~X4 = X6 => var(X6) = X5)) [true and false elimination 59]
59. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X4 & X0 = X2 & var(X6) = X3 & subst(X1,X2,X3) = X5) => ((~X4 = X6 & $true) => var(X6) = X5)) [rectify 25]
25. ! [X51,X58,X59,X60,X40,X61,X47] : ((X40 = X58 & X51 = X59 & var(X47) = X60 & subst(X58,X59,X60) = X61) => ((~X40 = X47 & $true) => var(X47) = X61)) [input]
170. ~tcheck(sK11,subst(sK12,sK13,sK14),sK15) (0:7) [cnf transformation 117]
117. var(sK9) = sK14 & sK9 != sK12 & tcheck(sK11,sK13,sK10) & tcheck(bind(sK12,sK10,sK11),sK14,sK15) & ~tcheck(sK11,subst(sK12,sK13,sK14),sK15) [skolemisation 99]
99. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 98]
98. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 69]
69. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
457. tcheck(sK11,sK14,sK15) {12} (3:4) [equality resolution 449]
449. some(sK15) != some(X3) | tcheck(sK11,sK14,X3) {12} (2:9) [forward demodulation 448,166]
166. var(sK9) = sK14 (0:4) [cnf transformation 117]
117. var(sK9) = sK14 & sK9 != sK12 & tcheck(sK11,sK13,sK10) & tcheck(bind(sK12,sK10,sK11),sK14,sK15) & ~tcheck(sK11,subst(sK12,sK13,sK14),sK15) [skolemisation 99]
99. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 98]
98. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 69]
69. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
448. some(sK15) != some(X3) | tcheck(sK11,var(sK9),X3) {12} (2:10) [superposition 146,440]
440. some(sK15) = lookup(sK9,sK11) {12} (1:6) [subsumption resolution 436,167]
167. sK9 != sK12 (0:3) [cnf transformation 117]
117. var(sK9) = sK14 & sK9 != sK12 & tcheck(sK11,sK13,sK10) & tcheck(bind(sK12,sK10,sK11),sK14,sK15) & ~tcheck(sK11,subst(sK12,sK13,sK14),sK15) [skolemisation 99]
99. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 98]
98. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 69]
69. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
436. some(sK15) = lookup(sK9,sK11) | sK9 = sK12 {12} (1:9) [superposition 434,192]
192. lookup(X5,X6) = lookup(X5,bind(X3,X0,X6)) | X3 = X5 (0:13) [equality resolution 191]
191. lookup(X5,X6) = lookup(X1,bind(X3,X0,X6)) | X3 = X5 | X1 != X5 (0:16) [equality resolution 190]
190. lookup(X1,X2) = lookup(X5,X6) | X3 = X5 | bind(X3,X0,X6) != X2 | X1 != X5 (0:19) [equality resolution 145]
145. lookup(X5,X6) = X4 | X3 = X5 | lookup(X1,X2) != X4 | bind(X3,X0,X6) != X2 | X1 != X5 (0:22) [cnf transformation 75]
75. ! [X0,X1,X2,X3,X4,X5,X6] : (X1 != X5 | bind(X3,X0,X6) != X2 | lookup(X1,X2) != X4 | X3 = X5 | lookup(X5,X6) = X4) [flattening 74]
74. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 != X5 | bind(X3,X0,X6) != X2 | lookup(X1,X2) != X4) | (X3 = X5 | lookup(X5,X6) = X4)) [ennf transformation 49]
49. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X5 & bind(X3,X0,X6) = X2 & lookup(X1,X2) = X4) => (X3 != X5 => lookup(X5,X6) = X4)) [flattening 48]
48. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X5 & bind(X3,X0,X6) = X2 & lookup(X1,X2) = X4) => (~X3 = X5 => lookup(X5,X6) = X4)) [true and false elimination 47]
47. ! [X0,X1,X2,X3,X4,X5,X6] : ((X1 = X5 & bind(X3,X0,X6) = X2 & lookup(X1,X2) = X4) => ((~X3 = X5 & $true) => lookup(X5,X6) = X4)) [rectify 16]
16. ! [X49,X45,X46,X47,X48,X40,X44] : ((X40 = X45 & bind(X47,X49,X44) = X46 & lookup(X45,X46) = X48) => ((~X40 = X47 & $true) => lookup(X40,X44) = X48)) [input]
434. some(sK15) = lookup(sK9,bind(sK12,sK10,sK11)) {12} (0:9) [sat splitting component]
146. some(X2) != lookup(X1,X0) | tcheck(X0,var(X1),X2) (0:11) [cnf transformation 76]
76. ! [X0,X1,X2] : (some(X2) != lookup(X1,X0) | tcheck(X0,var(X1),X2)) [ennf transformation 50]
50. ! [X0,X1,X2] : (some(X2) = lookup(X1,X0) => tcheck(X0,var(X1),X2)) [rectify 17]
17. ! [X44,X40,X50] : (some(X50) = lookup(X40,X44) => tcheck(X44,var(X40),X50)) [input]
384. sP0(bind(sK12,sK10,sK11),sK15,sK14) {10} (0:7) [sat splitting component]
434. some(sK15) = lookup(sK9,bind(sK12,sK10,sK11)) {12} (0:9) [sat splitting component]
432. some(sK15) = lookup(sK9,bind(sK12,sK10,sK11)) | sP0(bind(sK12,sK10,sK11),sK15,sK14) {8} (1:16) [forward demodulation 431,398]
398. sK5(bind(sK12,sK10,sK11),sK15,sK14) = sK9 {8} (2:9) [trivial inequality removal 395]
395. sK14 != sK14 | sK5(bind(sK12,sK10,sK11),sK15,sK14) = sK9 {8} (2:12) [superposition 217,382]
382. var(sK5(bind(sK12,sK10,sK11),sK15,sK14)) = sK14 {8} (0:10) [sat splitting component]
217. var(X0) != sK14 | sK9 = X0 (1:7) [superposition 118,166]
166. var(sK9) = sK14 (0:4) [cnf transformation 117]
117. var(sK9) = sK14 & sK9 != sK12 & tcheck(sK11,sK13,sK10) & tcheck(bind(sK12,sK10,sK11),sK14,sK15) & ~tcheck(sK11,subst(sK12,sK13,sK14),sK15) [skolemisation 99]
99. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 98]
98. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 69]
69. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
118. var(X0) != var(X1) | X0 = X1 (0:8) [cnf transformation 102]
102. ! [X0,X1] : ((var(X0) != var(X1) | X0 = X1) & (X0 != X1 | var(X0) = var(X1))) [nnf transformation 1]
1. ! [X0,X1] : (var(X0) = var(X1) <=> X0 = X1) [input]
431. sP0(bind(sK12,sK10,sK11),sK15,sK14) | some(sK15) = lookup(sK5(bind(sK12,sK10,sK11),sK15,sK14),bind(sK12,sK10,sK11)) {8} (1:22) [subsumption resolution 424,392]
392. app(X5,X6) != sK14 {8} (1:5) [superposition 139,382]
382. var(sK5(bind(sK12,sK10,sK11),sK15,sK14)) = sK14 {8} (0:10) [sat splitting component]
139. var(X0) != app(X1,X2) (0:6) [cnf transformation 39]
39. ! [X0,X1,X2] : var(X0) != app(X1,X2) [rectify 10]
10. ! [X28,X29,X30] : var(X28) != app(X29,X30) [input]
424. sP0(bind(sK12,sK10,sK11),sK15,sK14) | some(sK15) = lookup(sK5(bind(sK12,sK10,sK11),sK15,sK14),bind(sK12,sK10,sK11)) | app(sK6(bind(sK12,sK10,sK11),sK15,sK14),sK7(bind(sK12,sK10,sK11),sK15,sK14)) = sK14 (1:39) [resolution 153,169]
169. tcheck(bind(sK12,sK10,sK11),sK14,sK15) (0:7) [cnf transformation 117]
117. var(sK9) = sK14 & sK9 != sK12 & tcheck(sK11,sK13,sK10) & tcheck(bind(sK12,sK10,sK11),sK14,sK15) & ~tcheck(sK11,subst(sK12,sK13,sK14),sK15) [skolemisation 99]
99. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 98]
98. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 69]
69. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
153. ~tcheck(X2,X0,X1) | sP0(X2,X1,X0) | some(X1) = lookup(sK5(X2,X1,X0),X2) | app(sK6(X2,X1,X0),sK7(X2,X1,X0)) = X0 (0:28) [cnf transformation 116]
116. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | (var(sK5(X2,X1,X0)) = X0 & some(X1) = lookup(sK5(X2,X1,X0),X2)) | sP0(X2,X1,X0) | (app(sK6(X2,X1,X0),sK7(X2,X1,X0)) = X0 & tcheck(X2,sK6(X2,X1,X0),arrow(sK8(X2,X1,X0),X1)) & tcheck(X2,sK7(X2,X1,X0),sK8(X2,X1,X0)))) [skolemisation 115]
115. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X6] : (var(X6) = X0 & some(X1) = lookup(X6,X2)) | sP0(X2,X1,X0) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))) [rectify 101]
101. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | sP0(X2,X1,X0) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))) [definition folding 81,100]
100. ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9))) [predicate definition introduction]
81. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))) [flattening 80]
80. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | (? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))) [ennf transformation 53]
53. ! [X0,X1,X2] : (tcheck(X2,X0,X1) => (? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))) [rectify 20]
20. ! [X51,X50,X44] : (tcheck(X44,X51,X50) => (? [X40] : (var(X40) = X51 & some(X50) = lookup(X40,X44)) | ? [X40,X54,X55,X56] : (abs(X40,X55,X54) = X51 & arrow(X55,X56) = X50 & tcheck(bind(X40,X55,X44),X54,X56)) | ? [X53,X54,X52] : (app(X53,X54) = X51 & tcheck(X44,X53,arrow(X52,X50)) & tcheck(X44,X54,X52)))) [input]
1108. $false {10} (3:0) [trivial inequality removal 1107]
1107. sK14 != sK14 {10} (3:3) [superposition 1094,166]
166. var(sK9) = sK14 (0:4) [cnf transformation 117]
117. var(sK9) = sK14 & sK9 != sK12 & tcheck(sK11,sK13,sK10) & tcheck(bind(sK12,sK10,sK11),sK14,sK15) & ~tcheck(sK11,subst(sK12,sK13,sK14),sK15) [skolemisation 99]
99. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 98]
98. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 69]
69. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
1094. var(X18) != sK14 {10} (2:4) [superposition 138,459]
459. abs(sK1(sK14,sK15,bind(sK12,sK10,sK11)),sK3(sK14,sK15,bind(sK12,sK10,sK11)),sK2(sK14,sK15,bind(sK12,sK10,sK11))) = sK14 {10} (1:24) [resolution 384,149]
149. ~sP0(X0,X1,X2) | abs(sK1(X2,X1,X0),sK3(X2,X1,X0),sK2(X2,X1,X0)) = X2 (0:19) [cnf transformation 114]
114. ! [X0,X1,X2] : (~sP0(X0,X1,X2) | (abs(sK1(X2,X1,X0),sK3(X2,X1,X0),sK2(X2,X1,X0)) = X2 & arrow(sK3(X2,X1,X0),sK4(X2,X1,X0)) = X1 & tcheck(bind(sK1(X2,X1,X0),sK3(X2,X1,X0),X0),sK2(X2,X1,X0),sK4(X2,X1,X0)))) [skolemisation 113]
113. ! [X0,X1,X2] : (~sP0(X0,X1,X2) | ? [X3,X4,X5,X6] : (abs(X3,X5,X4) = X2 & arrow(X5,X6) = X1 & tcheck(bind(X3,X5,X0),X4,X6))) [rectify 112]
112. ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9))) [nnf transformation 100]
100. ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9))) [predicate definition introduction]
384. sP0(bind(sK12,sK10,sK11),sK15,sK14) {10} (0:7) [sat splitting component]
138. var(X0) != abs(X1,X2,X3) (0:7) [cnf transformation 38]
38. ! [X0,X1,X2,X3] : var(X0) != abs(X1,X2,X3) [rectify 9]
9. ! [X24,X25,X26,X27] : var(X24) != abs(X25,X26,X27) [input]
380. app(sK6(bind(sK12,sK10,sK11),sK15,sK14),sK7(bind(sK12,sK10,sK11),sK15,sK14)) = sK14 {6} (0:17) [sat splitting component]
382. var(sK5(bind(sK12,sK10,sK11),sK15,sK14)) = sK14 {8} (0:10) [sat splitting component]
384. sP0(bind(sK12,sK10,sK11),sK15,sK14) {10} (0:7) [sat splitting component]
368. sP0(bind(sK12,sK10,sK11),sK15,sK14) | var(sK5(bind(sK12,sK10,sK11),sK15,sK14)) = sK14 | app(sK6(bind(sK12,sK10,sK11),sK15,sK14),sK7(bind(sK12,sK10,sK11),sK15,sK14)) = sK14 (1:34) [resolution 152,169]
169. tcheck(bind(sK12,sK10,sK11),sK14,sK15) (0:7) [cnf transformation 117]
117. var(sK9) = sK14 & sK9 != sK12 & tcheck(sK11,sK13,sK10) & tcheck(bind(sK12,sK10,sK11),sK14,sK15) & ~tcheck(sK11,subst(sK12,sK13,sK14),sK15) [skolemisation 99]
99. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 98]
98. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 69]
69. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
152. ~tcheck(X2,X0,X1) | sP0(X2,X1,X0) | var(sK5(X2,X1,X0)) = X0 | app(sK6(X2,X1,X0),sK7(X2,X1,X0)) = X0 (0:26) [cnf transformation 116]
116. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | (var(sK5(X2,X1,X0)) = X0 & some(X1) = lookup(sK5(X2,X1,X0),X2)) | sP0(X2,X1,X0) | (app(sK6(X2,X1,X0),sK7(X2,X1,X0)) = X0 & tcheck(X2,sK6(X2,X1,X0),arrow(sK8(X2,X1,X0),X1)) & tcheck(X2,sK7(X2,X1,X0),sK8(X2,X1,X0)))) [skolemisation 115]
115. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X6] : (var(X6) = X0 & some(X1) = lookup(X6,X2)) | sP0(X2,X1,X0) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))) [rectify 101]
101. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | sP0(X2,X1,X0) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))) [definition folding 81,100]
100. ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9))) [predicate definition introduction]
81. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))) [flattening 80]
80. ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | (? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))) [ennf transformation 53]
53. ! [X0,X1,X2] : (tcheck(X2,X0,X1) => (? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))) [rectify 20]
20. ! [X51,X50,X44] : (tcheck(X44,X51,X50) => (? [X40] : (var(X40) = X51 & some(X50) = lookup(X40,X44)) | ? [X40,X54,X55,X56] : (abs(X40,X55,X54) = X51 & arrow(X55,X56) = X50 & tcheck(bind(X40,X55,X44),X54,X56)) | ? [X53,X54,X52] : (app(X53,X54) = X51 & tcheck(X44,X53,arrow(X52,X50)) & tcheck(X44,X54,X52)))) [input]
1154. $false {6} (2:0) [trivial inequality removal 1153]
1153. sK14 != sK14 {6} (2:3) [superposition 1137,166]
166. var(sK9) = sK14 (0:4) [cnf transformation 117]
117. var(sK9) = sK14 & sK9 != sK12 & tcheck(sK11,sK13,sK10) & tcheck(bind(sK12,sK10,sK11),sK14,sK15) & ~tcheck(sK11,subst(sK12,sK13,sK14),sK15) [skolemisation 99]
99. ? [X0,X1,X2,X3,X4,X5,X6] : (var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6) & ~tcheck(X2,subst(X3,X4,X5),X6)) [flattening 98]
98. ? [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) & ~tcheck(X2,subst(X3,X4,X5),X6)) [ennf transformation 69]
69. ~! [X0,X1,X2,X3,X4,X5,X6] : ((var(X0) = X5 & X0 != X3 & tcheck(X2,X4,X1) & tcheck(bind(X3,X1,X2),X5,X6)) => tcheck(X2,subst(X3,X4,X5),X6)) [rectify 30]
30. ~! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [negated conjecture 29]
29. ! [X47,X50,X44,X40,X51,X54,X56] : ((var(X47) = X54 & X40 != X47 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56)) [input]
1137. var(X8) != sK14 {6} (1:4) [superposition 139,380]
380. app(sK6(bind(sK12,sK10,sK11),sK15,sK14),sK7(bind(sK12,sK10,sK11),sK15,sK14)) = sK14 {6} (0:17) [sat splitting component]
139. var(X0) != app(X1,X2) (0:6) [cnf transformation 39]
39. ! [X0,X1,X2] : var(X0) != app(X1,X2) [rectify 10]
10. ! [X28,X29,X30] : var(X28) != app(X29,X30) [input]
------------------------------
Version: Vampire 3.0 (commit e3e1a68)
Termination reason: Refutation

Active clauses: 145
Passive clauses: 443
Generated clauses: 966
Final active clauses: 100
Final passive clauses: 220
Input formulas: 29
Initial clauses: 53

Duplicate literals: 1
Trivial inequalities: 6
Fw subsumption resolutions: 15
Fw demodulations: 89
Bw demodulations: 27

Equational tautologies: 21
Forward subsumptions: 346
Backward subsumptions: 16
Fw demodulations to eq. taut.: 30

Binary resolution: 67
Forward superposition: 242
Backward superposition: 394
Self superposition: 8
Equality resolution: 17

Split clauses: 8
Split components: 21
SAT solver clauses: 17
SAT solver unit clauses: 3
SAT solver binary clauses: 2
SAT solver learnt clauses: 2
SAT solver learnt literals: 3

Sat splits: 8
Sat splitting refutations: 6

Memory used [KB]: 1663
Time elapsed: 0.060 s
------------------------------
------------------------------
