% START OF SYSTEM OUTPUT
Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 599 next slice time: 3
lrs-1010_7_bs=unit_only:bsr=unit_only:cond=fast:drc=off:ep=on:flr=on:lcm=predicate:nwc=1.7:stl=30:sgo=on:spo=on:spl=sat:ssfq=1.1:ssnc=none:sp=occurrence_1 on SOT_yHnr
% (4030)Time limit reached!
% ------------------------------
% Version: Vampire 3.0 (revision 1803, CASC-24-only license, for any other Vampire license check http://vprover.org/)
% SZS status Timeout
% Termination phase: Saturation

% Active clauses: 429
% Passive clauses: 2805
% Generated clauses: 8039
% Final active clauses: 429
% Final passive clauses: 2376
% Input formulas: 33
% Initial clauses: 58
% 
% Duplicate literals: 229
% Fw subsumption resolutions: 77
% Fw demodulations: 320
% Condensations: 78
% 
% Equational tautologies: 15
% Forward subsumptions: 4530
% Fw demodulations to eq. taut.: 14
% 
% Binary resolution: 586
% Factoring: 18
% Forward superposition: 3531
% Backward superposition: 3164
% Self superposition: 12
% 
% Memory used [KB]: 6268
% Time elapsed: 0.400 s
% ------------------------------
% ------------------------------
% remaining time: 595 next slice time: 3
dis-1002_6_bs=off:drc=off:nwc=1.1:nicw=on:sd=1:ss=axioms:st=3.0:sio=off:spl=sat:sser=off_1 on SOT_yHnr
% (4032)Time limit reached!
% ------------------------------
% Version: Vampire 3.0 (revision 1803, CASC-24-only license, for any other Vampire license check http://vprover.org/)
% SZS status Timeout
% Termination phase: Saturation

% Active clauses: 359
% Passive clauses: 7160
% Generated clauses: 12677
% Final active clauses: 359
% Final passive clauses: 6673
% Input formulas: 33
% Initial clauses: 39
% 
% Selected by SInE selection: 23
% SInE iterations: 1
% 
% Duplicate literals: 273
% Trivial inequalities: 33
% Fw subsumption resolutions: 171
% Fw demodulations: 817
% 
% Simple tautologies: 50
% Equational tautologies: 11
% Forward subsumptions: 4272
% Fw demodulations to eq. taut.: 34
% 
% Binary resolution: 579
% Factoring: 1
% Forward superposition: 5943
% Backward superposition: 4820
% Self superposition: 40
% Equality factoring: 4
% Equality resolution: 40
% 
% Memory used [KB]: 5500
% Time elapsed: 0.400 s
% ------------------------------
% ------------------------------
% remaining time: 591 next slice time: 6
ott+1_5:1_bs=off:bsr=unit_only:cond=fast:drc=off:nwc=5:ss=axioms:sos=all_4 on SOT_yHnr
% (4033)Time limit reached!
% ------------------------------
% Version: Vampire 3.0 (revision 1803, CASC-24-only license, for any other Vampire license check http://vprover.org/)
% SZS status Timeout
% Termination phase: Saturation

% Active clauses: 400
% Passive clauses: 386
% Generated clauses: 613
% Final active clauses: 400
% Final passive clauses: 7
% Input formulas: 33
% Initial clauses: 31
% 
% Selected by SInE selection: 19
% SInE iterations: 3
% 
% Duplicate literals: 2
% 
% Equational tautologies: 227
% Forward subsumptions: 4
% 
% Binary resolution: 456
% Forward superposition: 148
% Backward superposition: 5
% 
% Memory used [KB]: 2430
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 584 next slice time: 28
ott+2_20_cond=fast:drc=off:flr=on:lcm=reverse:nwc=1.1:nicw=on:sio=off:spl=sat:sser=off:ssfp=10000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_25 on SOT_yHnr
% (4034)Time limit reached!
% ------------------------------
% Version: Vampire 3.0 (revision 1803, CASC-24-only license, for any other Vampire license check http://vprover.org/)
% SZS status Timeout
% Termination phase: Saturation

% Active clauses: 866
% Passive clauses: 14944
% Generated clauses: 43462
% Final active clauses: 854
% Final passive clauses: 12089
% Input formulas: 33
% Initial clauses: 56
% 
% Duplicate literals: 837
% Trivial inequalities: 32
% Fw subsumption resolutions: 217
% Fw demodulations: 2563
% Bw demodulations: 1
% Condensations: 1292
% 
% Simple tautologies: 34
% Equational tautologies: 34
% Forward subsumptions: 22852
% Backward subsumptions: 2000
% Fw demodulations to eq. taut.: 64
% 
% Binary resolution: 677
% Forward superposition: 20989
% Backward superposition: 17240
% Self superposition: 126
% Equality factoring: 3
% Equality resolution: 105
% 
% Memory used [KB]: 28400
% Time elapsed: 2.900 s
% ------------------------------
% ------------------------------
% remaining time: 555 next slice time: 3
lrs+1011_8:1_bd=off:bs=off:bsr=unit_only:cond=fast:drc=off:fde=none:nwc=1:nicw=on:stl=240:sos=all:sio=off:spo=on:ssac=none:ssfp=1000:ssfq=1.1:ssnc=all_dependent_1 on SOT_yHnr
% (4035)Time limit reached!
% ------------------------------
% Version: Vampire 3.0 (revision 1803, CASC-24-only license, for any other Vampire license check http://vprover.org/)
% SZS status Timeout
% Termination phase: Saturation

% Active clauses: 507
% Passive clauses: 4095
% Generated clauses: 5171
% Final active clauses: 507
% Final passive clauses: 3630
% Input formulas: 33
% Initial clauses: 56
% 
% Duplicate literals: 55
% Trivial inequalities: 43
% Fw subsumption resolutions: 180
% Fw demodulations: 47
% 
% Equational tautologies: 41
% Forward subsumptions: 564
% 
% Binary resolution: 3859
% Forward superposition: 483
% Backward superposition: 317
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 119
% Split components: 273
% Introduced splitting names: 77
% 
% Memory used [KB]: 5373
% Time elapsed: 0.400 s
% ------------------------------
% ------------------------------
% remaining time: 551 next slice time: 107
dis+1_2:3_bs=off:drc=off:lcm=predicate:nwc=3:nicw=on:ss=included:st=1.5:sos=on:sagn=off:sio=off:spl=sat:sser=off:sp=occurrence_97 on SOT_yHnr
% Refutation found. Thanks to Tanya!
% SZS status Theorem for SOT_yHnr
% SZS output start Proof for SOT_yHnr
fof(f779,plain,(
  $false),
  inference(subsumption_resolution,[],[f778,f130])).
fof(f130,plain,(
  ~tcheck(sK2,subst(sK3,sK4,sK5),sK6)),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  app(e1app,e2app) = sK5 & tcheck(sK2,sK4,sK1) & tcheck(bind(sK3,sK1,sK2),sK5,sK6) & ~tcheck(sK2,subst(sK3,sK4,sK5),sK6)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4,sK5,sK6])],[f76])).
fof(f76,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (app(e1app,e2app) = X4 & tcheck(X1,X3,X0) & tcheck(bind(X2,X0,X1),X4,X5) & ~tcheck(X1,subst(X2,X3,X4),X5))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((app(e1app,e2app) = X4 & tcheck(X1,X3,X0) & tcheck(bind(X2,X0,X1),X4,X5)) & ~tcheck(X1,subst(X2,X3,X4),X5))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : ((app(e1app,e2app) = X4 & tcheck(X1,X3,X0) & tcheck(bind(X2,X0,X1),X4,X5)) => tcheck(X1,subst(X2,X3,X4),X5))),
  inference(rectify,[],[f34])).
fof(f34,negated_conjecture,(
  ~! [X50,X44,X40,X51,X54,X56] : ((app(e1app,e2app) = X54 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56))),
  inference(negated_conjecture,[],[f33])).
fof(f33,conjecture,(
  ! [X50,X44,X40,X51,X54,X56] : ((app(e1app,e2app) = X54 & tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),X54,X56)) => tcheck(X44,subst(X40,X51,X54),X56))),
  file('/tmp/SystemOnTPTPFormReply3995/SOT_yHnrby','T-subst-app')).
fof(f778,plain,(
  tcheck(sK2,subst(sK3,sK4,sK5),sK6)),
  inference(forward_demodulation,[],[f777,f127])).
fof(f127,plain,(
  app(e1app,e2app) = sK5),
  inference(cnf_transformation,[],[f111])).
fof(f777,plain,(
  tcheck(sK2,subst(sK3,sK4,app(e1app,e2app)),sK6)),
  inference(forward_demodulation,[],[f776,f218])).
fof(f218,plain,(
  ( ! [X6,X4,X7,X5] : (app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X5,X6,app(X4,X7))) )),
  inference(equality_resolution,[],[f217])).
fof(f217,plain,(
  ( ! [X6,X4,X0,X7,X5] : (app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X6,app(X4,X7)) | X0 != X5) )),
  inference(equality_resolution,[],[f216])).
fof(f216,plain,(
  ( ! [X6,X4,X0,X7,X5,X1] : (app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,app(X4,X7)) | X1 != X6 | X0 != X5) )),
  inference(equality_resolution,[],[f215])).
fof(f215,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X1] : (app(subst(X5,X6,X4),subst(X5,X6,X7)) = subst(X0,X1,X2) | app(X4,X7) != X2 | X1 != X6 | X0 != X5) )),
  inference(equality_resolution,[],[f180])).
fof(f180,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3 | subst(X0,X1,X2) != X3 | app(X4,X7) != X2 | X1 != X6 | X0 != X5) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3 | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3)),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 != X5 | X1 != X6 | app(X4,X7) != X2 | subst(X0,X1,X2) != X3) | app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3)),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3)),
  inference(true_and_false_elimination,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((X0 = X5 & X1 = X6 & app(X4,X7) = X2 & subst(X0,X1,X2) = X3) => ($true => app(subst(X5,X6,X4),subst(X5,X6,X7)) = X3))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X62,X63,X64,X65,X53,X40,X51,X54] : ((X40 = X62 & X51 = X63 & app(X53,X54) = X64 & subst(X62,X63,X64) = X65) => ($true => app(subst(X40,X51,X53),subst(X40,X51,X54)) = X65))),
  file('/tmp/SystemOnTPTPFormReply3995/SOT_yHnrby',subst27)).
fof(f776,plain,(
  tcheck(sK2,app(subst(sK3,sK4,e1app),subst(sK3,sK4,e2app)),sK6)),
  inference(resolution,[],[f731,f764])).
fof(f764,plain,(
  tcheck(sK2,subst(sK3,sK4,e1app),arrow(sK14(bind(sK3,sK1,sK2),sK6,sK5),sK6))),
  inference(resolution,[],[f679,f128])).
fof(f128,plain,(
  tcheck(sK2,sK4,sK1)),
  inference(cnf_transformation,[],[f111])).
fof(f679,plain,(
  ( ! [X0] : (~tcheck(sK2,X0,sK1) | tcheck(sK2,subst(sK3,X0,e1app),arrow(sK14(bind(sK3,sK1,sK2),sK6,sK5),sK6))) )),
  inference(resolution,[],[f668,f162])).
fof(f162,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~tcheck(bind(X2,X0,X1),e1app,X4) | tcheck(X1,subst(X2,X3,e1app),X4) | ~tcheck(X1,X3,X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3,X4] : (~tcheck(X1,X3,X0) | ~tcheck(bind(X2,X0,X1),e1app,X4) | tcheck(X1,subst(X2,X3,e1app),X4))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3,X4] : ((~tcheck(X1,X3,X0) | ~tcheck(bind(X2,X0,X1),e1app,X4)) | tcheck(X1,subst(X2,X3,e1app),X4))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3,X4] : ((tcheck(X1,X3,X0) & tcheck(bind(X2,X0,X1),e1app,X4)) => tcheck(X1,subst(X2,X3,e1app),X4))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X50,X44,X40,X51,X56] : ((tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),e1app,X56)) => tcheck(X44,subst(X40,X51,e1app),X56))),
  file('/tmp/SystemOnTPTPFormReply3995/SOT_yHnrby','T-subst-IH-app1')).
fof(f668,plain,(
  tcheck(bind(sK3,sK1,sK2),e1app,arrow(sK14(bind(sK3,sK1,sK2),sK6,sK5),sK6))),
  inference(subsumption_resolution,[],[f665,f625])).
fof(f625,plain,(
  ( ! [X10,X11,X9] : (abs(X9,X10,X11) != sK5) )),
  inference(superposition,[],[f160,f619])).
fof(f619,plain,(
  app(sK12(bind(sK3,sK1,sK2),sK6,sK5),sK13(bind(sK3,sK1,sK2),sK6,sK5)) = sK5),
  inference(trivial_inequality_removal,[],[f616])).
fof(f616,plain,(
  sK5 != sK5 | app(sK12(bind(sK3,sK1,sK2),sK6,sK5),sK13(bind(sK3,sK1,sK2),sK6,sK5)) = sK5),
  inference(superposition,[],[f605,f127])).
fof(f605,plain,(
  ( ! [X2,X1] : (app(X1,X2) != sK5 | app(sK12(bind(sK3,sK1,sK2),sK6,sK5),sK13(bind(sK3,sK1,sK2),sK6,sK5)) = sK5) )),
  inference(superposition,[],[f160,f308])).
fof(f308,plain,(
  abs(sK7(sK5,sK6,bind(sK3,sK1,sK2)),sK9(sK5,sK6,bind(sK3,sK1,sK2)),sK8(sK5,sK6,bind(sK3,sK1,sK2))) = sK5 | app(sK12(bind(sK3,sK1,sK2),sK6,sK5),sK13(bind(sK3,sK1,sK2),sK6,sK5)) = sK5),
  inference(resolution,[],[f250,f142])).
fof(f142,plain,(
  ( ! [X2,X0,X1] : (~sP0(X0,X1,X2) | abs(sK7(X2,X1,X0),sK9(X2,X1,X0),sK8(X2,X1,X0)) = X2) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : (~sP0(X0,X1,X2) | (abs(sK7(X2,X1,X0),sK9(X2,X1,X0),sK8(X2,X1,X0)) = X2 & arrow(sK9(X2,X1,X0),sK10(X2,X1,X0)) = X1 & tcheck(bind(sK7(X2,X1,X0),sK9(X2,X1,X0),X0),sK8(X2,X1,X0),sK10(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9,sK10])],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : (~sP0(X0,X1,X2) | ? [X3,X4,X5,X6] : (abs(X3,X5,X4) = X2 & arrow(X5,X6) = X1 & tcheck(bind(X3,X5,X0),X4,X6)))),
  inference(rectify,[],[f114])).
fof(f114,plain,(
  ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X2,X1,X0] : (~sP0(X2,X1,X0) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f250,plain,(
  sP0(bind(sK3,sK1,sK2),sK6,sK5) | app(sK12(bind(sK3,sK1,sK2),sK6,sK5),sK13(bind(sK3,sK1,sK2),sK6,sK5)) = sK5),
  inference(subsumption_resolution,[],[f242,f240])).
fof(f240,plain,(
  ( ! [X11] : (var(X11) != sK5) )),
  inference(superposition,[],[f140,f127])).
fof(f140,plain,(
  ( ! [X2,X0,X1] : (var(X0) != app(X1,X2)) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2] : var(X0) != app(X1,X2)),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X28,X29,X30] : var(X28) != app(X29,X30)),
  file('/tmp/SystemOnTPTPFormReply3995/SOT_yHnrby','DIFF-var-app')).
fof(f242,plain,(
  sP0(bind(sK3,sK1,sK2),sK6,sK5) | var(sK11(bind(sK3,sK1,sK2),sK6,sK5)) = sK5 | app(sK12(bind(sK3,sK1,sK2),sK6,sK5),sK13(bind(sK3,sK1,sK2),sK6,sK5)) = sK5),
  inference(resolution,[],[f129,f145])).
fof(f145,plain,(
  ( ! [X2,X0,X1] : (~tcheck(X2,X0,X1) | sP0(X2,X1,X0) | var(sK11(X2,X1,X0)) = X0 | app(sK12(X2,X1,X0),sK13(X2,X1,X0)) = X0) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | (var(sK11(X2,X1,X0)) = X0 & some(X1) = lookup(sK11(X2,X1,X0),X2)) | sP0(X2,X1,X0) | (app(sK12(X2,X1,X0),sK13(X2,X1,X0)) = X0 & tcheck(X2,sK12(X2,X1,X0),arrow(sK14(X2,X1,X0),X1)) & tcheck(X2,sK13(X2,X1,X0),sK14(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12,sK13,sK14])],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X6] : (var(X6) = X0 & some(X1) = lookup(X6,X2)) | sP0(X2,X1,X0) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))),
  inference(rectify,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | sP0(X2,X1,X0) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))),
  inference(definition_folding,[],[f78,f109])).
fof(f78,plain,(
  ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | ? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5)))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2] : (~tcheck(X2,X0,X1) | (? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2] : (tcheck(X2,X0,X1) => (? [X10] : (var(X10) = X0 & some(X1) = lookup(X10,X2)) | ? [X6,X7,X8,X9] : (abs(X6,X8,X7) = X0 & arrow(X8,X9) = X1 & tcheck(bind(X6,X8,X2),X7,X9)) | ? [X3,X4,X5] : (app(X3,X4) = X0 & tcheck(X2,X3,arrow(X5,X1)) & tcheck(X2,X4,X5))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X51,X50,X44] : (tcheck(X44,X51,X50) => (? [X40] : (var(X40) = X51 & some(X50) = lookup(X40,X44)) | ? [X40,X54,X55,X56] : (abs(X40,X55,X54) = X51 & arrow(X55,X56) = X50 & tcheck(bind(X40,X55,X44),X54,X56)) | ? [X53,X54,X52] : (app(X53,X54) = X51 & tcheck(X44,X53,arrow(X52,X50)) & tcheck(X44,X54,X52))))),
  file('/tmp/SystemOnTPTPFormReply3995/SOT_yHnrby','T-inv')).
fof(f129,plain,(
  tcheck(bind(sK3,sK1,sK2),sK5,sK6)),
  inference(cnf_transformation,[],[f111])).
fof(f160,plain,(
  ( ! [X4,X2,X0,X3,X1] : (abs(X0,X1,X2) != app(X3,X4)) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3,X4] : abs(X0,X1,X2) != app(X3,X4)),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X31,X32,X33,X34,X35] : abs(X31,X32,X33) != app(X34,X35)),
  file('/tmp/SystemOnTPTPFormReply3995/SOT_yHnrby','DIFF-abs-app')).
fof(f665,plain,(
  tcheck(bind(sK3,sK1,sK2),e1app,arrow(sK14(bind(sK3,sK1,sK2),sK6,sK5),sK6)) | abs(sK7(sK5,sK6,bind(sK3,sK1,sK2)),sK9(sK5,sK6,bind(sK3,sK1,sK2)),sK8(sK5,sK6,bind(sK3,sK1,sK2))) = sK5),
  inference(resolution,[],[f659,f142])).
fof(f659,plain,(
  sP0(bind(sK3,sK1,sK2),sK6,sK5) | tcheck(bind(sK3,sK1,sK2),e1app,arrow(sK14(bind(sK3,sK1,sK2),sK6,sK5),sK6))),
  inference(backward_demodulation,[],[f641,f251])).
fof(f251,plain,(
  sP0(bind(sK3,sK1,sK2),sK6,sK5) | tcheck(bind(sK3,sK1,sK2),sK12(bind(sK3,sK1,sK2),sK6,sK5),arrow(sK14(bind(sK3,sK1,sK2),sK6,sK5),sK6))),
  inference(subsumption_resolution,[],[f244,f240])).
fof(f244,plain,(
  sP0(bind(sK3,sK1,sK2),sK6,sK5) | var(sK11(bind(sK3,sK1,sK2),sK6,sK5)) = sK5 | tcheck(bind(sK3,sK1,sK2),sK12(bind(sK3,sK1,sK2),sK6,sK5),arrow(sK14(bind(sK3,sK1,sK2),sK6,sK5),sK6))),
  inference(resolution,[],[f129,f147])).
fof(f147,plain,(
  ( ! [X2,X0,X1] : (~tcheck(X2,X0,X1) | sP0(X2,X1,X0) | var(sK11(X2,X1,X0)) = X0 | tcheck(X2,sK12(X2,X1,X0),arrow(sK14(X2,X1,X0),X1))) )),
  inference(cnf_transformation,[],[f118])).
fof(f641,plain,(
  e1app = sK12(bind(sK3,sK1,sK2),sK6,sK5)),
  inference(trivial_inequality_removal,[],[f626])).
fof(f626,plain,(
  sK5 != sK5 | e1app = sK12(bind(sK3,sK1,sK2),sK6,sK5)),
  inference(superposition,[],[f236,f619])).
fof(f236,plain,(
  ( ! [X4,X3] : (app(X3,X4) != sK5 | e1app = X3) )),
  inference(superposition,[],[f157,f127])).
fof(f157,plain,(
  ( ! [X2,X0,X3,X1] : (app(X0,X1) != app(X2,X3) | X0 = X2) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((app(X0,X1) != app(X2,X3) | (X0 = X2 & X1 = X3)) & (X0 != X2 | X1 != X3 | app(X0,X1) = app(X2,X3)))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((app(X0,X1) != app(X2,X3) | (X0 = X2 & X1 = X3)) & ((X0 != X2 | X1 != X3) | app(X0,X1) = app(X2,X3)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : (app(X0,X1) = app(X2,X3) <=> (X0 = X2 & X1 = X3))),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X8,X9,X10,X11] : (app(X8,X9) = app(X10,X11) <=> (X8 = X10 & X9 = X11))),
  file('/tmp/SystemOnTPTPFormReply3995/SOT_yHnrby','EQ-app')).
fof(f731,plain,(
  ( ! [X2,X3] : (~tcheck(sK2,X2,arrow(sK14(bind(sK3,sK1,sK2),sK6,sK5),X3)) | tcheck(sK2,app(X2,subst(sK3,sK4,e2app)),X3)) )),
  inference(resolution,[],[f723,f165])).
fof(f165,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~tcheck(X1,X3,X0) | tcheck(X1,app(X2,X3),X4) | ~tcheck(X1,X2,arrow(X0,X4))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3,X4] : (~tcheck(X1,X2,arrow(X0,X4)) | ~tcheck(X1,X3,X0) | tcheck(X1,app(X2,X3),X4))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3,X4] : ((~tcheck(X1,X2,arrow(X0,X4)) | ~tcheck(X1,X3,X0)) | tcheck(X1,app(X2,X3),X4))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3,X4] : ((tcheck(X1,X2,arrow(X0,X4)) & tcheck(X1,X3,X0)) => tcheck(X1,app(X2,X3),X4))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X52,X44,X53,X54,X50] : ((tcheck(X44,X53,arrow(X52,X50)) & tcheck(X44,X54,X52)) => tcheck(X44,app(X53,X54),X50))),
  file('/tmp/SystemOnTPTPFormReply3995/SOT_yHnrby','T-app')).
fof(f723,plain,(
  tcheck(sK2,subst(sK3,sK4,e2app),sK14(bind(sK3,sK1,sK2),sK6,sK5))),
  inference(resolution,[],[f669,f128])).
fof(f669,plain,(
  ( ! [X0] : (~tcheck(sK2,X0,sK1) | tcheck(sK2,subst(sK3,X0,e2app),sK14(bind(sK3,sK1,sK2),sK6,sK5))) )),
  inference(resolution,[],[f657,f163])).
fof(f163,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~tcheck(bind(X2,X0,X1),e2app,X4) | tcheck(X1,subst(X2,X3,e2app),X4) | ~tcheck(X1,X3,X0)) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3,X4] : (~tcheck(X1,X3,X0) | ~tcheck(bind(X2,X0,X1),e2app,X4) | tcheck(X1,subst(X2,X3,e2app),X4))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3,X4] : ((~tcheck(X1,X3,X0) | ~tcheck(bind(X2,X0,X1),e2app,X4)) | tcheck(X1,subst(X2,X3,e2app),X4))),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3,X4] : ((tcheck(X1,X3,X0) & tcheck(bind(X2,X0,X1),e2app,X4)) => tcheck(X1,subst(X2,X3,e2app),X4))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X50,X44,X40,X51,X56] : ((tcheck(X44,X51,X50) & tcheck(bind(X40,X50,X44),e2app,X56)) => tcheck(X44,subst(X40,X51,e2app),X56))),
  file('/tmp/SystemOnTPTPFormReply3995/SOT_yHnrby','T-subst-IH-app2')).
fof(f657,plain,(
  tcheck(bind(sK3,sK1,sK2),e2app,sK14(bind(sK3,sK1,sK2),sK6,sK5))),
  inference(subsumption_resolution,[],[f653,f625])).
fof(f653,plain,(
  tcheck(bind(sK3,sK1,sK2),e2app,sK14(bind(sK3,sK1,sK2),sK6,sK5)) | abs(sK7(sK5,sK6,bind(sK3,sK1,sK2)),sK9(sK5,sK6,bind(sK3,sK1,sK2)),sK8(sK5,sK6,bind(sK3,sK1,sK2))) = sK5),
  inference(backward_demodulation,[],[f640,f316])).
fof(f316,plain,(
  tcheck(bind(sK3,sK1,sK2),sK13(bind(sK3,sK1,sK2),sK6,sK5),sK14(bind(sK3,sK1,sK2),sK6,sK5)) | abs(sK7(sK5,sK6,bind(sK3,sK1,sK2)),sK9(sK5,sK6,bind(sK3,sK1,sK2)),sK8(sK5,sK6,bind(sK3,sK1,sK2))) = sK5),
  inference(resolution,[],[f252,f142])).
fof(f252,plain,(
  sP0(bind(sK3,sK1,sK2),sK6,sK5) | tcheck(bind(sK3,sK1,sK2),sK13(bind(sK3,sK1,sK2),sK6,sK5),sK14(bind(sK3,sK1,sK2),sK6,sK5))),
  inference(subsumption_resolution,[],[f246,f240])).
fof(f246,plain,(
  sP0(bind(sK3,sK1,sK2),sK6,sK5) | var(sK11(bind(sK3,sK1,sK2),sK6,sK5)) = sK5 | tcheck(bind(sK3,sK1,sK2),sK13(bind(sK3,sK1,sK2),sK6,sK5),sK14(bind(sK3,sK1,sK2),sK6,sK5))),
  inference(resolution,[],[f129,f149])).
fof(f149,plain,(
  ( ! [X2,X0,X1] : (~tcheck(X2,X0,X1) | sP0(X2,X1,X0) | var(sK11(X2,X1,X0)) = X0 | tcheck(X2,sK13(X2,X1,X0),sK14(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f118])).
fof(f640,plain,(
  e2app = sK13(bind(sK3,sK1,sK2),sK6,sK5)),
  inference(trivial_inequality_removal,[],[f627])).
fof(f627,plain,(
  sK5 != sK5 | e2app = sK13(bind(sK3,sK1,sK2),sK6,sK5)),
  inference(superposition,[],[f238,f619])).
fof(f238,plain,(
  ( ! [X8,X7] : (app(X7,X8) != sK5 | e2app = X8) )),
  inference(superposition,[],[f158,f127])).
fof(f158,plain,(
  ( ! [X2,X0,X3,X1] : (app(X0,X1) != app(X2,X3) | X1 = X3) )),
  inference(cnf_transformation,[],[f122])).
% SZS output end Proof for SOT_yHnr
% ------------------------------
% Version: Vampire 3.0 (revision 1803, CASC-24-only license, for any other Vampire license check http://vprover.org/)
% SZS status Unsatisfiable

% Active clauses: 156
% Passive clauses: 452
% Generated clauses: 544
% Final active clauses: 142
% Final passive clauses: 301
% Input formulas: 33
% Initial clauses: 56
% 
% Selected by SInE selection: 33
% SInE iterations: 1
% 
% Duplicate literals: 5
% Trivial inequalities: 9
% Fw subsumption resolutions: 35
% Fw demodulations: 19
% Bw demodulations: 9
% 
% Equational tautologies: 10
% Forward subsumptions: 60
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 306
% Forward superposition: 38
% Backward superposition: 117
% Equality resolution: 2
% 
% Memory used [KB]: 511
% Time elapsed: 0.024 s
% ------------------------------
% ------------------------------
% Success in time 4.832 s

% END OF SYSTEM OUTPUT
% RESULT: SOT_yHnrby - Vampire---3.0 says Theorem - CPU = 4.41 WC = 4.83 
% OUTPUT: SOT_yHnrby - Vampire---3.0 says Refutation - CPU = 4.41 WC = 4.83 

